Exam 2 is cumulative, although the focus is on Loops, Methods, Classes, and File Paths.

The format is the same as Exam 1. There are two parts, worth 16.5% of your course grade:

- 13.5% in class
- 3% take-home


## In-Class Portion

2/27 Tues (come early).
Bring:<br>
- Pencils/Pen
- Student ID

Closed notes. No computers.<br>
Somewhat more code writing than Exam 1.<br>
75 minutes.

One version of last term's exam [is here](COGS18_Sec_B_Exam_2_Fa23.pdf) if you want to use it to study.


## Take-Home Portion

Distributed on Datahub.<br>
~3 Coding exercises, similar to assignment questions.<br>
Open notes. Open Google.<br>
Work alone. No other humans. No ChatGPT or other AIs.<br>
Submit on Datahub before class 2/29 10:59am. (Just like an assignment, but you cannot talk to others or ChatGPT.)

## Preparation

Thursday in class will be a review session (Feb 22).

This practice exam imitates the take-home portion, but is much longer.

# COGS 18 - Practice Exam 2

This is the second practice exam for Fall 2023. It covers topics through File Paths in the Command Line lecture.

This is worth 0 points and worth 0% of your grade, but the real total exam is out of 16.5 points, worth 16.5% of your grade. 

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.** (Note that you can add additional cells, if you want to test things out.

<span style="color:red">Note: There are WAYYYYYY more questions on this practice exam than on the real exam. This is because a number of students have asked for additional practice problems. So, do not be intimidated by the length here, but feel free to use it for extra practice!</span>

## The Rules (for the real exam)

**These rules do _not_ apply to this practice exam, but are a copy of what is expected for the real take-home exam.** Therefore, please do ask people for help with this practice!

- Complete this exam _on your own_.
- This is open-notes & open-Internet.
- You may _not_ talk to any humans about this exam. 
- You may _not_ talk to any AIs about this exam.
- **The following are all *prohibited*:**
    - text/phone/online chat communication
    - posting questions to a message board where a human could respond (Campuswire, Discord, Chegg, any similar site)
    - viewing exam questions from a message board (as described above)
    - asking anyone via any form about a question on this test
    - asking any AI in any form about a question on this test
    - searching for answers using text from these questions directly 
- **Instructional staff will *not* help you.**
    - Do _not_ ask about about the exam during office hours, not even clarification questions. We will not respond.
    - If an exam question needs clarification, instead ask in a _private_ post on Piazza. At our discrection, we may clarify _publicly_ for all to benefit. You may not get a quick response. If you ask on the last day, you may not get a response. You probably have the correct interpretation.
    - We will _not_ help you towards an answer. You can consult your notes, the lectures, and Google search on topics, but we will not response to questions to get you unstuck.
    - If you have a technical issue completing the exam, message (Piazza or email) Dr. Hempel ASAP.

 <span style="color: red;">Note: </span> There _is_ a chance for partial credit on some questions, so _some_ code is better than _no_ code. Even if it throws an error, having some code that partially answers the question will benefit you/your grade.

## Part 1: Methods, Debugging, and Code Style

### Q1 - Method I

`isalpha()` is a built-in string method that determines if all the characters in a string are in the alphabet (meaning letters or numbers: A-Z, a-z).

Use the `isalpha()` method on two strings (of your choosing) below to create two variables:
- `true_var` | should store the value `True`
- `false_var` | should store the value `False`

<!-- grade_id: cell-09260f9c394a5905, points: 0, grade: false, solution: true -->
```python
# your code here
```

<!-- grade_id: cell-e9a5e70357e98aef, points: 0.2, grade: true, solution: false -->
```python
assert true_var is not None
# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert true_var
```

<!-- grade_id: cell-12001d8195362930, points: 0.2, grade: true, solution: false -->
```python
assert false_var is not None
# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert not false_var
```

### Q2 - Method II

We've provided a variable called `prof_day` below. Generate a dictionary `prof_events`, whose keys are each of the unique elements in `prof_day` and whose corresponding values are the number of times each value shows up in `prof_day`.
 
Notes:
- `prof_day` has been defined for you below
- `.count()` is a list method that counts the number of appearances of a specified element in a list.
- **Do not hard-code**. Use your code to process the list so that we could change the list and you'd still generate the correct answer automatically. That is DON'T write something like:
```python
prof_events = {'write one exam':2}
```
because if we changed the list this answer would be wrong.


```python
# these variables provided for you
# you will not be able to edit this cell
prof_day = ['write one exam', 'write one exam', 'grade exam', 'answer email', 'office hours', 'staff meeting', 'answer email']

```

<!-- grade_id: cell-6992bbe86b5d62a3, points: 0, grade: true, solution: false -->
```python
# hidden test being used in grading
# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# this is the kind of thing we do to check you didn't hard code
# we change one value in list you were originally shown 
# then the tests go on to execute below

prof_day = ['write one exam', 'write one exam', 'answer email', 'answer email', 'office hours', 'staff meeting', 'answer email']
```

<!-- grade_id: cell-43c896bbf4a8a127, points: 0, grade: false, solution: true -->
```python
# your code here
```

<!-- grade_id: cell-7cc63f7c144c54af, points: 0.2, grade: true, solution: false -->
```python
assert prof_events is not None

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# check that keys are all what they should be
assert all(elem in prof_events.keys() for elem in ['write one exam', 'answer email', 'office hours', 'staff meeting'])
```

<!-- grade_id: cell-1de53862761ee0f0, points: 0.2, grade: true, solution: false -->
```python
# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# check that values are what they should be
assert sum(prof_events.values()) == 7
assert prof_events == {
    'write one exam': 2,
    'answer email': 3,
    'office hours': 1,
    'staff meeting': 1,
}
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q3 - Debugging Loops

In the cell below, there is a function that is not operating as anticipated.

`count_consonants` is _supposed_ to take a string as input and count the number of each consonant in the `input_string` (input parameter), storing the output in a dictionary where the key is the consonant and the value is the number of times it appears in the input. 

Note that we want capitalization to be considered such that the capital and lower case letters are counted together, and the upper case letter is used as the key (i.e. 'A' and 'a' would both be counted under the key 'A')

Specifically, if the `input_string` were "Mohammed", the function ould return: `{'M': 3, 'H': 1, 'D': 1}`

Debug, edit, and test the function provided below so that it accomplishes the intended goal.

```python
# This is the function that is not operating as anticipated and that has bad code style.
# Debug, edit, and test the function provided below so that it accomplishes the intended goal.
def count_consonants():
    out = {}
    vowels = ['A','E','I','O','U']

    for char in input_string:
        if char not in vowels:
            if char not in out:
                out[char] = 1

# You can copy and paste this code (not the cell) to start your coding
```

<!-- grade_id: cell-bf67b586499664bc, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q6-test1, points: 0.35, grade: true, solution: false -->
```python
assert count_consonants is not None
```

<!-- grade_id: q6-test2, points: 0.35, grade: true, solution: false -->
```python
assert callable(count_consonants)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert count_consonants('AaBbCc') == {'B': 2, 'C': 2}
assert count_consonants('Taylor') == {'T': 1, 'Y': 1, 'L': 1, 'R': 1}
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q4 - Debugging Loops

In the cell below, there is a function that is not operating as anticipated.

`new_encrypt` is _supposed_ to take a string as input and return the reverse alphabetical string as output, meaning, given a character, return the __lowercase__ character in the reverse position within the alphabet. 

For example, if the input character is 'B' or 'b', the function would store 'y' in the output string. If the character is 'A' or 'a', the function would store 'z'. (and vice versa: the input 'Z' or 'z' would store 'a').

Specifically, if the `input_string` were "AbZ", the function would `return`: 'zya'

Debug, edit, and test the function provided below so that it accomplishes the intended goal.

```python
# This is the function that is not operating as anticipated and has bad code style.
# Debug, edit, and test the function provided below so that it accomplishes the intended goal.
def new_encrypt():
    alpha = 'abcdefghijklmnopqrstuvwxyz'
    reverse_alpha = 'zyxwvutsrqponmlkjihgfedcba'

    for char in input_string:

        if char in alpha:
            position = alpha.find(char)
            output_string += reverse_alpha[position]
        else:
            output_string = None
            break
            
# You can copy and paste the code (not the cell) to start your debugging
```

<!-- grade_id: q7-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q7-test1, points: 0.25, grade: true, solution: false -->
```python
assert new_encrypt
```

<!-- grade_id: q7-test2, points: 0.25, grade: true, solution: false -->
```python
assert callable(new_encrypt)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert new_encrypt(input_string='abc') == 'zyx'
assert new_encrypt('MOHAMMED') == 'nlsznnvw'

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q5 - Debugging (1.5 points)

The following not-totally-functioning `give_ticket()` function has been provided for you. The *goal* of this function, is given two inputs (`speed` and `zone`), the function should return the boolean True if the person should receive a ticket and the boolean False otherwise.

The rules for giving a ticket depend on the `zone` which can take the following values/rules:
- `private`&nbsp;| ticket if  `speed` is above 5
- `local`&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;|   ticket if `speed` is above 25
- `highway` |  ticket if `speed` is above 65


For example, if the driver's `speed` was 25 and their zone was 'local', the function would return `False`. However, if in the same zone, and their speed were 26, the function would return `True`. 

Consider the function provided below and then debug (you can copy and paste the code below) to accomplish the task specified above. 

Note: Do not change the name of the function (`give_ticket`), and the parameter names provided in the instructions (`speed`, `zone`) must be used.

```python
# This is the no-functioning function
# Debug and edit it to accomplish the task specified above.
def give_ticket(self, speed, zone):

    if zone = 'private' or speed > 5:
        print(True)
    elif zone = 'local' or speed > 25:
        print(True)
    elif zone = 'highway' or speed > 65:
        print(True)
    else:
        print(False)
    
# You can copy and paste the code (not the cell) to start your debugging
```

<!-- grade_id: cell-ee5cbe431e55aeb8, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-d3ecf602a2d3eee9, points: 0.25, grade: true, solution: false -->
```python
assert give_ticket

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert give_ticket(speed=6, zone='private') 
assert give_ticket(speed=26, zone='local') 
assert give_ticket(speed=66, zone='highway') 
```

<!-- grade_id: cell-f70941db79b59474, points: 0.25, grade: true, solution: false -->
```python
assert callable(give_ticket)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert not give_ticket(speed=5, zone='private') 
assert not give_ticket(speed=10, zone='local') 
assert not give_ticket(speed=45, zone='highway') 
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


## Part 2: Loops & Functions

Note this section has *more* questions than your real exam will have. This is just to get you additional practice.

### Q6 - Function execution

The cell below includes a function that has been provided for you. After running the cell below (to define this function), execute (use) this function to create the following variables under the specified conditions:

- `mulitply_default` | execute the `multiply_all` function using the default parameter such that the output will carry out the function on a list of input values containing the integers 2, 3, and 4. 
- `multiplication_5` | Use the same input values as above to obtain an answer from `multiply_all` function, but this time using multiplication_factor as 5.
- `my_output` &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;| Execute the `multiply_all` function such that `my_output` will store a list with 4 values `[10,10,10,10]` , such that the multiplication_factor = -5.  

```python
def multiply_all(my_list, multiplication_factor=2):
    
    output_list = []
    for val in my_list:
        new_val = val * multiplication_factor
        output_list.append(new_val)
    
    return output_list
```

<!-- grade_id: q4-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL BELOW
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: q4-test1, points: 0.25, grade: true, solution: false -->
```python
assert mulitply_default is not None

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert  mulitply_default == [4,6,8]
```

<!-- grade_id: q4-test2, points: 0.25, grade: true, solution: false -->
```python
assert multiplication_5 is not None

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert multiplication_5  == [10, 15, 20]
```

<!-- grade_id: q4-test3, points: 0.25, grade: true, solution: false -->
```python
assert my_output is not None

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert my_output  == [10,10,10,10]
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q7 - Function Execution

The function `count_sentences()` has been provided for you below. Note that `split()` is a string method that splits a string at the character specified within the parentheses of the method call.

If this isn't clear to you, you might want to try the following in a cell and see what happens:
```python
'milk, butter, eggs'.split(',')
```

**Part 1** 

First, in the cell provided below containing `YOUR ANSWER HERE`, describe what the `count_sentences()` function accomplishes. This should include both 1) the overall task accomplished and 2) how the specific code used accomplishes this task.

```python
def count_sentences(text):
    sentences = text.split('.')
    
    if sentences[-1] == '':
        num_sentences = len(sentences) - 1
    else:
        num_sentences = len(sentences)
        
    return num_sentences
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL BELOW
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-f060a0b5dbbc43db, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 2**

Execute the `count_sentences()` function by passing it a string that will cause it to return a value of `2`; assign this output to the variable `out_2`.

<!-- grade_id: cell-b60e1b7443ba1cb1, points: 0, grade: false, solution: true -->
```python
# your code here
```

<!-- grade_id: cell-b92e934f9e45a69a, points: 0, grade: true, solution: false -->
```python
assert out_2
# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert out_2 == 2
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q8 - `count_int` 

Write a function `count_int` that will take a list as the input and count the number of integers in that list, returning this value from the function. 

For example, if the list had 5 elements, with 2 of the elements being strings and 3 of the elements were integers, then the function would return 3.

Note: You can assume the input to this function will be a list and do *not* need to write code to check whether or not this is true.

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q5-test1, points: 0.25, grade: true, solution: false -->
```python
assert callable(count_int)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert count_int((1,2,3,4,5)) == 5
assert count_int([1,3,'string', 4.4, None, True ]) == 2
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q9 - `while` loop, indexing into multiple lists, list `append`

Write a function `zip_max` that takes in two lists as input.

`zip_max` should return a new list that is the maximum item of the corresponding items between the lists, for example:

```python
nums1 = [1,  20, 30, 4,  5]
nums2 = [50, 4,  3,  20, 10, 0]

assert zip_max(nums1, nums2) == [50, 20, 30, 20, 10]
```

Note the input lists may not be the same length.

Use at least `while` loop and list `.append()` inside the loop to accomplish this task.

**Hints:**

- Set an index counter before the loop (name it whatever you want).
- Set an empty output list before the loop (name it whatever you want).
- The `while` loop should keep looping so long as the index counter is less than the length of both lists.
- In the loop, compare the appropriate items in the lists and append the appropriate one to the output list.
- Don't forget to increment the index counter in the loop.

<!-- grade_id: cell-fdd3939da79c3975, points: 0, grade: false, solution: true -->
```python
def zip_max(lst1, lst2):
    # your code here
```

<!-- grade_id: cell-0cca5022c0814421, points: 0, grade: true, solution: false -->
```python
nums1 = [1,  20, 30, 4,  5]
nums2 = [50, 4,  3,  20, 10, 0]

assert zip_max(nums1, nums2) == [50, 20, 30, 20, 10]


nums1 = [10, 30, 4,  5, 0]
nums2 = [-1, -1]

assert zip_max(nums1, nums2) == [10, 30]
```

### Q10 - `modify_string`


In this question, you're going to define a function called `modify_string()` that will change `input_string` using the key-value pairs in `new_language` (provided below), ultimately returning the updated name from the function.

Specifically, if any of the letters in your name (`input_string`) are keys in the dictionary `new_language` below, use that key's value in that letter's place in the string being output from the function. If the letter is not in `new_language`, add the original character from `input_string` to the string being output from the function.

For example, if `input_string` were 'Shannon', `modif_string()` would `return` 'Shӓnnðn' (the a and the o have the modified character from `new_language`)


This function should have a single parameter: `input_string` - the string to be modified.

```python
new_language = {'a': 'ӓ',
                'e' : 'ɚ',
                'i' : '¡',
                'o' : 'ð',
                'u' : 'û',
                'A' : 'Ӓ',
                'E' : 'ɛ',
                'O' : 'Ó',
                'U' : 'Ü'}
```

<!-- grade_id: q5-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: test1, points: 0.25, grade: true, solution: false -->
```python
assert modify_string

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
assert modify_string('taylor') == 'tӓylðr'
```

<!-- grade_id: test2, points: 0.25, grade: true, solution: false -->
```python
assert callable(modify_string)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you
# check functioning for capital letters
assert modify_string('AaBbCcEe') == 'ӒӓBbCcɛɚ'

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q11 - `calculate_points`

Write a function `calculate_points` that has a single input parameter: `dice`. `dice` can be assumed to be a list of string values, where each string corresponds to a value corresponding to each face of dice (1,2,3,4,5,6)

This function should calculate the total number of "pointer" faces of a dice, where ***a "pointer" is either a 4, 5, or 6***. Each "pointer" is worth one point. All other faces are worth zero points. The function should `return` the total number of points in the `dice` as an integer.

For example, if your `dice` input was: 

```python
['1', '3', '4', '5', '6']
```

executing the `calculate_points()` function with the above list as the input to the `dice` parameter, the function would `return` the integer 3 (one point each for the `'4'`, `'5'`, and `'6'`.

If there are no pointers in the input `dice`, the function should return 0.

<!-- grade_id: q6-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q6-testa, points: 0.25, grade: true, solution: false -->
```python
assert callable(calculate_points)

```

<!-- grade_id: q6-testb, points: 0.25, grade: true, solution: false -->
```python
# hidden tests for function above

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# assure returns zero if no pointers
assert calculate_points(['1', '2', '3']) == 0
assert calculate_points(['4','5']) == 2
assert calculate_points([5, 6]) == 0 # these are ints, we said strings
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)



### Q12 - `to_taxes` - Looping and Dictionaries

Here you will write a function that, given a salary as input, will `return` how much of that person's salary is going toward taxes.

To accomplish this, write a function `to_taxes()` that takes two parameters as input: `salary` and `tax_brackets`. The default value for the `tax_brackets` parameter should be the `default_brackets` dictionary provided below. Then, use code constructs discussed in class to accomplish the above goal.

Note that the dictionary `default_brackets` provided will have to be used to accomplish this goal. The keys in `default_brackets` are the proportion of an individual's salary that will go toward taxes if their salary is within the range in a given key's value (tuple). For example, an individual who makes \$10,000 would contribute 12\% of their salary in taxes, given the information in `tax_brackets` (specifically becasue 10000 falls between 9876 and 40125).

Accordingly, `to_taxes(salary=10000)` should `return` 1200.0 (which corresponds to 12\% of the input $10,000 salary).

Dictionaries in Python are ordered (by insertion order), so iterating over `tax_brackets` will process the keys in the expected order.

```python
default_brackets = {
    0.10 : (0, 11000),
    0.12 : (11001, 44725),
    0.22 : (44726, 95375),
    0.24 : (95376, 182100),
    0.32 : (182101, 231250),
    0.35 : (231251, 578125)
}
```

<!-- grade_id: cell-f48b7b2362162c98, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-46841edd2405b87e, points: 0, grade: true, solution: false -->
```python
assert to_taxes

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert callable(to_taxes)
assert to_taxes(10000) == 1000.0
assert to_taxes(182101) == 58272.32
```

<!-- grade_id: cell-48691fd539be2c35, points: 0, grade: true, solution: false -->
```python
# hidden test to check function accomplishes task

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# check for 22% bracket
# check that values on edges included
assert to_taxes(44726) == 9839.72
assert to_taxes(95375) == 20982.5
```

<!-- grade_id: cell-11839f4b8385d811, points: 0, grade: true, solution: false -->
```python
# hidden test to check default parameter correct

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# check for a different tax_brackets
assert to_taxes(100, tax_brackets={0.5: (0,1000)}) == 50.0
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q13 - Loops and Dictionary Methods

Write a function `shared_values` that takes two dictionaries as input.

The function should return two new dictionaries (in a tuple) such that the new dictionaries contain the entries where some value is shared between the two dictionaries.

That is, each dictionary should be filtered down to only those entries whose value appears as a value somewhere in the other dictionary.

For example:

```python
d1 = {
    'a': 1,
    'b': 2,
    'c': 3,
    'd': 4,
}

d2 = {
    'Z': 6,
    'Y': 5,
    'X': 4,
    'W': 3,
}

# Only the values 3 and 4 are shared between the dictionaries.

out1 = {'c': 3, 'd': 4}
out2 = {'X': 4, 'W': 3}

assert shared_values(d1, d2) == (out1, out2)
```

**Hints:**
- You may want to use some of the [dictionary methods](https://www.w3schools.com/python/python_ref_dictionary.asp).
- You may need two separate (non-nested) loops, one to process each dictionary.

<!-- grade_id: cell-9b0c62c3966a82f4, points: 0, grade: false, solution: true -->
```python
def shared_values(dict1, dict2):
    out1 = {}
    out2 = {}
    
    # your code here
    
    return (out1, out2)
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you may test your function here.
```

<!-- grade_id: cell-512d4527dace1b8b, points: 0, grade: true, solution: false -->
```python
d1 = {
    'a': 1,
    'b': 2,
    'c': 3,
    'd': 4,
}

d2 = {
    'Z': 6,
    'Y': 5,
    'X': 4,
    'W': 3,
}

d3 = {
    '😺' : 3,
    '🌧️' : 6
}

assert shared_values(d1, d2) == ({'c': 3, 'd': 4}, {'X': 4, 'W': 3})

assert shared_values(d1, d3) == ({'c': 3}, {'😺' : 3})

assert shared_values(d3, d3) == (d3, d3)
```

### Q14 Filtering and Transforming A List With Loops and `.append()`

Write a function `odds_only_doubled` that takes in one list as input.

It should return an output list that:

- Retains the items in the input list that are odd numbers (discarding the others), _and_
- doubles those numbers.

For example:

```python
in_list = [1, 2, 3, 3, 4, 0]

assert odds_only_doubled(in_list) == [2, 6, 6]
```

**Hints:**
- Recall the condition `x % 2 == 1` will check if `x` is odd.

<!-- grade_id: cell-fcc930f7ab913422, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you may test your function here.
```

<!-- grade_id: cell-fa395030ed89ae59, points: 0, grade: true, solution: false -->
```python
in_list = [1, 2, 3, 3, 4, 0]

assert odds_only_doubled(in_list) == [2, 6, 6]

assert odds_only_doubled([]) == []

assert odds_only_doubled([5, 0, 7, 0]) == [10, 14]
```

## Part 3: Classes

Note this section has *more* questions than your real exam will have. This is where students tend to struggle the most as the material is new and puts all the pieces we've talked about so far this quarter together, so I wanted you to have extra practice.

### Q15 - `Movie`

For this question, imagine you want to create a movie game with a handful of characters. You want this game to have a movie theme.

Here, define a class `Movie`.

This class should have two instance attributes: `movie_name` and `actor`, which will be strings specifying the movie_name and an actor in that movie.

This class will also have a method `introduce()`. This method should `return` (*not* just `print`) a string, similar to:

```
'Inception starring Leonardo DiCaprio.' 
```

...where 'Inception' corresponds to whatever is stored in `Movie()` object's `movie_name` attribute is and 'Leonardo DiCaprio' corresponds to whatever is stored in the `Movie()` object's `actor` attribute.

Note that you can make the introduction say whatever you want as long as the output is a string mentioning both the movie name and the actor.

<!-- grade_id: q9-solution, points: 0, grade: false, solution: true -->
```python
# your code here

```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q9-test1, points: 0.5, grade: true, solution: false -->
```python
# Tests for the object
assert Movie

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

my_movie = Movie('Arrival', 'Amy Adams')
assert isinstance(my_movie, Movie)
assert my_movie.movie_name == 'Arrival'
assert my_movie.actor == 'Amy Adams'
```

<!-- grade_id: q9-test2, points: 0.5, grade: true, solution: false -->
```python
# Tests for introduce method
assert callable(Movie)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert callable(my_movie.introduce)
assert isinstance(my_movie.introduce(), str)
assert my_movie.actor in my_movie.introduce()
assert my_movie.movie_name in my_movie.introduce()

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q16 - `NewYear` - Classes and Dictionary Iteration

For this question, let's create a class to celebrate the Lunar New Year!

Here, define a class `NewYear`.

This class should have:

- **one class attribute**: `zodiac_signs` that stores the dictionary specified in the cell below. This dictionary contains the zodiac sign as the key and the birth years that correspond to that sign.
- **one instance attribute**:  `year` that takes someone's birth `year` as input upon creation of a `NewYear` instance.

This class will also have a method `return_sign()`. This method should `return` (*not* just `print`) a string, similar to:

```
'You were born in the year of the Dragon!' 
```

...where 'Dragon' corresponds to the key from the `zodiac_signs` corresponding to the value stored in the `year` attribute. (Be sure captitalization and punctuation match the string specified here. The only thing that should change is the last word in the string.)

For example `NewYear(year=2021).return_sign()` should return 'You were born in the year of the Ox!', since 2021 is a value in the `zodiac_signs` dictionary for the key 'Ox'. 

```python
# provided for you
zodiac_signs = {
    'Ox' : [1937, 1949, 1961, 1973, 1985, 1997, 2009, 2021],
    'Tiger' : [1938, 1950, 1962, 1974, 1986, 1998, 2010, 2022], 
    'Rabbit' : [1939, 1951, 1963, 1975, 1987, 1999, 2011, 2023], 
    'Dragon' : [1940, 1952, 1964, 1976, 1988, 2000, 2012, 2024], 
    'Snake' : [1941, 1953, 1965, 1977, 1989, 2001, 2013, 2025], 
    'Horse' : [1942, 1954, 1966, 1978, 1990, 2002, 2014, 2026], 
    'Goat/Sheep' : [1943, 1955, 1967, 1979, 1991, 2003, 2015, 2027], 
    'Monkey' : [1944, 1956, 1968, 1980, 1992, 2004, 2016, 2028], 
    'Rooster' : [1945, 1957, 1969, 1981, 1993, 2005, 2017, 2029], 
    'Dog' : [1946, 1958, 1970, 1982, 1994, 2006, 2018, 2030], 
    'Pig' : [1947, 1959, 1971, 1983, 1995, 2007, 2019, 2031], 
    'Rat' : [1936, 1948, 1960, 1972, 1984, 1996, 2008, 2020]
} 
```

<!-- grade_id: q10-solution, points: 0, grade: false, solution: true -->
```python
# your code here

```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your function (optional)
```

<!-- grade_id: q10-test1, points: 0.25, grade: true, solution: false -->
```python
# Tests for the object
assert NewYear

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

dragon = NewYear(1988)
assert isinstance(dragon, NewYear)
```

<!-- grade_id: q10-test2, points: 0.25, grade: true, solution: false -->
```python
assert callable(NewYear)
# Tests for classs attribute

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert dragon.year ==  1988
```

<!-- grade_id: q10-test3, points: 0.25, grade: true, solution: false -->
```python
# Tests for instance attribute

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert isinstance(dragon.zodiac_signs, dict)
assert 'Ox' in dragon.zodiac_signs.keys()
```

<!-- grade_id: q10-test4, points: 0.5, grade: true, solution: false -->
```python
# tests for method
# be sure punctuation matches instructions

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert dragon.return_sign() == 'You were born in the year of the Dragon!'
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q17 - `BasketballGame`

Generate a class called `BasketballGame`.

This class should have four instance attributes: `home_team`, `away_team`, `home_points`, and `away_points`

It should also have one method, `play_game()`. 

Within the `play_game` method, add code that would determine who the winner of the game is (determined by the team with the most points), returning 'Winner: Team Name Here' but with the winning team's name there.

For ties, return: 'Winner: tie'

<!-- grade_id: q11-solution, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# use this cell to test/execute/check your class (optional)
```

<!-- grade_id: q11-test1, points: 0, grade: true, solution: false -->
```python
assert callable(BasketballGame)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

game = BasketballGame(home_team='Warriors', away_team='Sixers', 
                      home_points=100, away_points=105)

# test attributes
assert game.home_team == 'Warriors'
assert game.away_team == 'Sixers'
assert game.home_points == 100
assert game.away_points == 105
```

<!-- grade_id: q11-test2, points: 0, grade: true, solution: false -->
```python
assert callable(BasketballGame.play_game)

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

# test method
assert game.play_game() == 'Winner: Sixers'
```

This "blank" cell included intentionally. Do not do anything here. (It's being used in grading.)


### Q18 - `ClassRoster`

Define a class `ClassRoster` that meets the following specifications:

- has two instance attributes:
    - `students`; initialized as an empty list
    - `course`; which the user specifies when creating a `ClassRoster()` object
- a single method `add_student()`, which will have `pid` and `name` as arguments, and will add these two inputs as a single-entry dictionary with `pid` as the key and the student's `name` as the value to the `students` list each time the method is called

<!-- grade_id: cell-d4d9cea6256f62c0, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-22422a3990743b88, points: 0, grade: true, solution: false -->
```python
assert callable(ClassRoster)
my_course = ClassRoster('COGS 18')


# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert my_course.course == 'COGS 18'
```

<!-- grade_id: cell-bd0c067e35d23349, points: 0, grade: false, solution: true -->
```python
# your code here
```

<!-- grade_id: cell-c250ea29b92bbfb1, points: 0, grade: false, solution: true -->
```python
# your code here
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


### Q19 - `ShoppingList` 

Define a class `ShoppingList()` to keep track of your to do list.

This should have a single instance attribute: `shoppinglist`, which is initialized as an empty list.

It will then have two methods: `add_item` and `remove_item`.

`add_item()`:
- will have two parameters `item`, and `importance` (`importance` takes the default value `True`)
- if `importance` is `True`:
    the string in `item` will be added to the top of `shoppinglist`:
- if `importance` is not true:
    the string it `item` will be added to the end of `shoppinglist`

`remove_item()`:
- will have a single parameter `item`
- will remove the item specified from `shoppinglist`

Note: there is a list method `insert()`, which operates in place to add a value to a list at the index specified with the general syntax `my_list.insert(index, item_to_add)`.

<!-- grade_id: cell-e696ac9734d6920a, points: 0, grade: false, solution: true -->
```python
# your code here
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-b43d68904a417d44, points: 0, grade: true, solution: false -->
```python
assert callable(ShoppingList)
my_shoppinglist = ShoppingList()

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

assert my_shoppinglist.shoppinglist == []
```

<!-- grade_id: cell-16d48156f776f357, points: 0, grade: true, solution: false -->
```python
# check add_item method

# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

my_shoppinglist.add_item('A')
my_shoppinglist.add_item('B')
assert my_shoppinglist.shoppinglist == ['B','A']
```

<!-- grade_id: cell-90ca75a4a66f62cc, points: 0, grade: true, solution: false -->
```python
# check add_item with non-default value


# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

my_shoppinglist.add_item('C', importance=False)
assert my_shoppinglist.shoppinglist == ['B','A', 'C']
```

<!-- grade_id: cell-41ae8cf0c3d717c5, points: 0, grade: true, solution: false -->
```python
# check remove_item


# below this would normally be hidden tests,
# but because this is a practice exam 
# we are going to show it to you

my_shoppinglist.remove_item('A')
assert my_shoppinglist.shoppinglist == ['B', 'C']
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


## Part 4: File Paths

### Q20 - File Paths

Consider the following directory structure:

- **`/`**
    - **`home/`**
        - **`alice/`**
            - `analysis.py`
            - `data.csv`
            - `make_plot.py`
            - `plot.png`
        - **`drhempel/`**
            - **`cogs18/`**
                - **`assignments/`**
                    - `A1.ipynb`
                    - `A2.ipynb`
                    - `A3.ipynb`
                - **`exams/`**
                    - `E1_Take_Home.ipynb`
                    - `E1_In_Class.tex`
                    - `E2_Take_Home.ipynb`
                    - `E2_In_Class.tex`
                - **`scratch/`**
                    - `scratch.py`
                    - `helpers.py`


We are interacting as the user `drhempel`. And here is where we are, according to `pwd`:

```
$ pwd
/home/drhempel/cogs18
```

**Part 1:** We are currently in the `cogs18` directory. What is the _absolute_ path to `A1.ipynb`?

<!-- grade_id: cell-3a0ea2432983a3c0, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 2:** We are currently in the `cogs18` directory. What is the _relative_ path to `A1.ipynb`?

<!-- grade_id: cell-1981c831ab53af36, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 3:** We are currently in the `cogs18` directory. What is the path to `A1.ipynb` using `~`?

<!-- grade_id: cell-e8175fe7ca17985f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 4:** We are currently in the `cogs18` directory. What is the path to where we are using `~`?

<!-- grade_id: cell-937b0b49a4d08a4c, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 5:** We are currently in the `cogs18` directory. What is the _relative_ path to Alice's `make_plot.py`?

<!-- grade_id: cell-e9d4094803131f04, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 6:** Say we are now in the `scratch` directory. What do I put in the blank to run the `scratch.py` script?

```
$ pwd
/home/drhempel/cogs18/scratch
$ python ____________
```

<!-- grade_id: cell-e9f6726b73670d2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Part 7:** We are still in the `scratch` directory. What could I put in the blank to run Alice's `analysis.py` script? (There are at least three possible answers.)

```
$ pwd
/home/drhempel/cogs18/scratch
$ python _______________________
```

<!-- grade_id: cell-7b317da14f161245, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

(Note that, in practice, you usually cannot see or run the files in other user's home folders, but we didn't talk about "file permissions" so we are ignoring that.)

### (Practice) Exam II complete!

Good work - you're done with the second (practice) exam!

A reminder that for the real thing you will need to save and submit on datahub. We will not have your exam until you click submit and see this show up under 'submitted assignments'.