# Assignment 1: Getting Started

This assignment covers Variables, Operators & Functions.

This assignment is out of 8 points, worth 8% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**

## How to complete assignments

Whenever you see:

```python
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (meaning, delete) these lines of code with some code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. **These should "pass silently" (give no output), letting you know that you're on the right track.**

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

### Printing Variables

A reminder that you can (and should) print and check variables as you go.  Note that this is for your benefit, not Python's.  It allows you to see the values that are stored within your variables, but it does nothing for the rest of the program.

This allows you to check what values they hold, and fix things if anything unexpected happens.

```python
# Define a variable
math_result = 2 * 4

# Print out the value(s) of a variable.
print(math_result)
```

Output:
```
8
```

### Restarting the Kernel

A reminder that sometimes if you run cells out of order, you can end up overwriting things in your namespace. 

If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.

Even if everything seems to be working, it's a nice check to 'Restart & Run All' to make sure everything runs properly in order.

## Example Question

Define a variable called `my_int` that stores the value 17. 

```python
# YOUR CODE HERE
my_int = 17
```

```python
# This is a test cell. It checks that the code you wrote does what it's supposed to
assert  my_int                     # This tests that a variable exists
assert  type(my_int) == int        # This tests that a variable is of a particular type
assert  my_int == 17               # This tests that a variable has a specified value
```

## Part I - Defining Variables

First, we will start by defining some Python variables. 

### Q1 - Create Variables (1.25 points)

Create the following variables:
- a variable called `number` that stores the value 929
- a variable called `decimal` that stores the value 9.29
- a variable called `word` that stores the value 'cogs 18' (note: there is a space between cogs and 18)
- a variable called `truth` that stores the value True
- a variable called `nada` that stores the value None

<!-- grade_id: cell-aba100abfe23c75f, points: 0, grade: false, solution: true -->
```python
# Create the variables

### BEGIN SOLUTION
number = 929
decimal = 9.29
word = 'cogs 18'
truth = True
nada = None
### END SOLUTION
```

<!-- grade_id: cell-6d742eec19ccc8b1, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `number` is an int, and has the correct value
assert type(number) == int
assert number == 929
```

<!-- grade_id: cell-4757efea652f25f7, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `decimal` is a float, and has the correct value
assert type(decimal) == float
assert decimal == 9.29
```

<!-- grade_id: cell-8b0f7e2f80328a33, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `words` is a string, and has the correct value
assert type(word) == str
assert word == 'cogs 18'
```

<!-- grade_id: cell-bea89cf2bc700120, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `truth` is a boolean, and has the correct value
assert type(truth) == bool
assert truth == True
```

<!-- grade_id: cell-a9284fa7620f0807, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q1

# Check that variable `nada` is NoneType, and has the correct value
assert type(nada) == type(None)
assert nada == None
```

## Part II - Operations

### Q2 - Python as a Calculator (1 point)

Reference the variable names provided below (`int_1`, `int_2`, `float_1`, `float_2`) to carry out the following calculations:

- Set `ans_1` as the result of `int_1` divided by `int_2`
- Set `ans_2` as the result of `float_1` multiplied by `float_2`, with `int_2` then added afterward
- Set `ans_3` as the result of `float_1` divided by `int_1`, the result of which is then raised to the power 2.
- Set `ans_4` as the result of the remainder of `int_1` divided by `int_2`

Note: your answer should use the *variables* provided (i.e. `int_1`), *not* the values they store (i.e. `27`).

```python
# These variables provided for you
int_1 = 27
int_2 = 9

float_1 = 9.5
float_2 = 19.0
```

<!-- grade_id: cell-470235f72bd8d68d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ans_1 = int_1 / int_2
ans_2 = float_1 * float_2 + int_2
ans_3 = (float_1 / int_1) ** 2
ans_4 = int_1 % int_2
### END SOLUTION
```

<!-- grade_id: cell-12af4e4d3179f69e, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert type(ans_1) == float
### BEGIN HIDDEN TESTS
import math 
assert math.isclose(ans_1, 3.0, rel_tol = 0.01)
### END HIDDEN TESTS
```

<!-- grade_id: cell-aeeac9dbb5198fef, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert type(ans_2) == float
### BEGIN HIDDEN TESTS
assert math.isclose(ans_2, 189.5, rel_tol = 0.01)
### END HIDDEN TESTS
```

<!-- grade_id: cell-4bdf00498b0db006, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert type(ans_3) == float
### BEGIN HIDDEN TESTS
assert math.isclose(ans_3, 0.123, rel_tol = 0.01)
### END HIDDEN TESTS
```

<!-- grade_id: cell-9df0702d77cd4bf2, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q2  [Note: Includes Hidden Tests]

assert type(ans_4) == int
### BEGIN HIDDEN TESTS
assert ans_4 == 0
### END HIDDEN TESTS
```

### Q3 - String Manipulations (0.5 points)

Using the strings provided, concatenate them together into the string `Python is fun!`.Save this result to a variable called `str_fun`.

Do the same, but concatenate to generate the string `Python is really fun!`. Save this result to a variable `str_really_fun`.

Note: make sure to add spaces between words when concatenating together. 

```python
# These variables provided for you
str_1 = 'Python'
str_2 = 'is'
str_3 = 'really'
str_4 = 'fun!'
```

<!-- grade_id: cell-eff2d0b53b5430eb, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
str_fun = str_1 + ' ' + str_2 + ' ' + str_4
str_really_fun = str_1 + ' ' + str_2 + ' ' + str_3 + ' ' + str_4
### END SOLUTION
```

<!-- grade_id: cell-b9a1407c39acc29e, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q3

assert type(str_fun) == str
```

<!-- grade_id: cell-b297ff5f4ce9a659, points: 0.25, grade: true, solution: false -->
```python
##### TESTS FOR Q3 [Note: Includes Hidden Tests]

assert str_fun == 'Python is fun!'
### BEGIN HIDDEN TESTS
assert str_really_fun == 'Python is really fun!'
### END HIDDEN TESTS
```

### Q4 - Boolean Comparisons (0.5 points)

Using the variables provided, do the following boolean logical comparisons:
- Define `bool_comp_1` to be the logical `or` of `bool_1` and `bool_2`
- Define `bool_comp_2` to be the logical `and` of  `bool_1` and `bool_2`

```python
# These variables provided for you
bool_1 = True
bool_2 = False
```

<!-- grade_id: cell-bf6144adc1f1e7f7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bool_comp_1 = bool_1 or bool_2
bool_comp_2 = bool_1 and bool_2
### END SOLUTION
```

<!-- grade_id: cell-eeb7dcc8687131df, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q4 [Note: Includes Hidden Tests]
assert type(bool_comp_1) == bool

### BEGIN HIDDEN TESTS
assert bool_comp_1
### END HIDDEN TESTS
```

<!-- grade_id: cell-7a05eded0276311c, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q4 [Note: Includes Hidden Tests]
assert type(bool_comp_2) == bool

### BEGIN HIDDEN TESTS
assert not bool_comp_2
### END HIDDEN TESTS
```

### Q5 - Value Comparisons (0.5 points)

Do the following value comparisons:
- Use comparison operators to determine if `comp_val_1` is exactly equal to the value stored in `comp_val_2`
    - Store the result of this comparison to a variable called `comp_1`


- Use comparison operators to determine if `comp_val_4` is less than or equal to the value stored in `comp_val_3`
    - Store the result of this comparison to a variable called `comp_2`

```python
# These variables provided for you
comp_val_1 = 13
comp_val_2 = 22
comp_val_3 = 13
comp_val_4 = 11
```

<!-- grade_id: cell-41f5d60aa2e09728, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
comp_1 = comp_val_1 == comp_val_2
comp_2 = comp_val_4 <= comp_val_3
### END SOLUTION
```

<!-- grade_id: cell-313f6320a6e8c70a, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q5 [Note: Includes Hidden Tests]
assert type(comp_1) == bool

### BEGIN HIDDEN TESTS
assert not comp_1
### END HIDDEN TESTS
```

<!-- grade_id: cell-283d8c591109241d, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q5 [Note: Includes Hidden Tests]
assert type(comp_2) == bool

### BEGIN HIDDEN TESTS
assert comp_2
### END HIDDEN TESTS
```

## Part III - Functions

This part covers the basics of writing and using functions.

### Function Example

The following is a function example provided to you, so you can remind yourself of the structure of functions.

You don't need to do anything with this function but understand it.

```python
# define the function
def basic_function(input_number):
    
    # Multiply the input by 2
    output = input_number * 2 
    
    return output
```

```python
# execute `basic_function`
print(basic_function(2))
print(basic_function(5))
```

Output:
```
4
10
```

### Q6 - Calculator (1 point)

In Q2, we saw that Python can act as a calculator; however, those standalone statements are not particularly reusable when you want to do calculations with a different set of numbers.

Here, we will make a series of functions to do all the mathematical operations that we want to do. 

Define four functions, called `summation`, `subtraction`, `multiplication`, and `division`. 

Each of them will take two inputs `num1` and `num2` (that could be `int` or `float`), and use them as follows:
- `summation` will add `num1` to `num2` and return the result
- `subtraction` will carry out subtraction: `num1` minus `num2` and return the result
- `multiplication` will multiply `num1` by `num2` and return the result
- `division` will divide `num1` by `num2` and return the result

<!-- grade_id: cell-6da5de076983020f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def summation(num1, num2):
    return num1 + num2

def subtraction(num1, num2):
    return num1 - num2

def multiplication(num1, num2):
    return num1 * num2

def division(num1, num2):
    return num1 / num2
### END SOLUTION
```

```python
# Cell included for you to check your function works as intended
```

<!-- grade_id: cell-2db5b9469ddd3f5e, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q6 [Note: Includes Hidden Tests]

assert callable(summation)
assert type(summation(2, 2)) == int
assert type(summation(2.5, 1)) == float
assert summation(1, 1) == 2

### BEGIN HIDDEN TESTS
assert summation(10, 5) == 15
### END HIDDEN TESTS
```

<!-- grade_id: cell-0e21e32dc2acc953, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q6 [Note: Includes Hidden Tests]

assert callable(subtraction)
assert type(subtraction(2, 2)) == int
assert type(summation(2.5, 0.5)) == float
assert subtraction(1, 1) == 0

### BEGIN HIDDEN TESTS
assert subtraction(16, 7) == 9
### END HIDDEN TESTS
```

<!-- grade_id: cell-89646145e140818b, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q6 [Note: Includes Hidden Tests]

assert callable(multiplication)
assert type(multiplication(2, 2)) == int
assert type(multiplication(2.5, 2)) == float
assert multiplication(1, 1) == 1

### BEGIN HIDDEN TESTS
assert multiplication(18, 2) == 36
### END HIDDEN TESTS
```

<!-- grade_id: cell-065fccebdd8cda04, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q6 [Note: Includes Hidden Tests]

assert callable(division)
assert type(division(2, 2)) == float
assert type(division(2.5, 0.5)) == float
assert division(1, 1) == 1

### BEGIN HIDDEN TESTS
assert division(24, 3) == 8
### END HIDDEN TESTS
```

### Q7 - Extending our Calculator  (0.5 points)

Let's now add some more functions to our calculator: 

- Write a function called `squared` that takes one parameter, called `num`, and returns the square of it. 
- Write a function called `square_root` that takes one parameter, called `num`, and returns the square root of it.

Reminder: you can get the square root of a number by raising it to the power 1/2. You do not need to `import` any modules/packages that we have not yet discussed in class to answer this question.

<!-- grade_id: cell-249d571080483d7e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def squared(num):
    return num ** 2

def square_root(num):
    return num ** (1/2)
### END SOLUTION
```

```python
# Cell included for you to check your function works as intended
```

<!-- grade_id: cell-d9c2b8004a926b36, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q7 [Note: Includes Hidden Tests]

assert callable(squared)
assert type(squared(2)) == int
assert squared(5) == 25

### BEGIN HIDDEN TESTS
assert squared(7) == 49
### END HIDDEN TESTS
```

<!-- grade_id: cell-259646bd6dbbc035, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q7 [Note: Includes Hidden Tests]

assert callable(square_root)
assert type(square_root(4)) == float
assert square_root(144) == 12

### BEGIN HIDDEN TESTS
assert square_root(49) == 7
### END HIDDEN TESTS
```

### Example: Chaining Functions Together

The goal of using functions is partly to be able to re-use code segments. 

It's also a way to be able to flexibly apply and combine different functions together, to build up to more complex and interesting programs. 

The cell below shows a couple examples of using our calculator functions that we have now defined, in combination with each other. 

```python
# Add the multiplication of 2 and 3 to the multiplication of 3 and 4
number1 = multiplication(2, 3)
number2 = multiplication(3, 4)
output1 = summation(number1, number2)
print('Output1: \t', output1)

print('')

# Note that we can use the outputs of functions directly as inputs to other functions
#   Therefore, the above construction is equivalent to the following:
output2 = summation(multiplication(2, 3), multiplication(3, 4))
print('Output2: \t', output2)

# When we have something like above, the functions that are inputs get evaluated first:
#   So, first, `multiplication(2, 3)` evaluates as 6, and `multiplication(3, 4)` evaluates as 12
#   This then reduces to calling `summation(6, 12)`, which returns our answer of 18
```

Output:
```
Output1: 	 18

Output2: 	 18
```

### Q8 - Using Our Calculator  (1 point)

Let's now *use our new calculator functions* to do some calculations!

- Add 2 squared to the square root of 49
    - Save the result to a variable called `calc_a`
- Subtract the division of 12 by 3 from the the multiplication of 2 and 3
    - Save the result to a variable called `calc_b`
- Multiply the square root of 36 by the sum of 7 and 6
    - Save the result to a variable called `calc_c`
- Divide the square of 5 by the square root of 100
    - Save the result to a variable called `calc_d`
    
Note: The answers here should utilize the functions you created previously.

<!-- grade_id: cell-78166757e331fd3f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
calc_a = summation(squared(2), square_root(49))
calc_b = subtraction(multiplication(2, 3), division(12, 3))
calc_c = multiplication(square_root(36), summation(7, 6))
calc_d = division(squared(5), square_root(100))

print(calc_a, calc_b, calc_c, calc_d)
### END SOLUTION
```

Output:
```
11.0 2.0 78.0 2.5
```

```python
# Cell included for you to check your function works as intended
```

<!-- grade_id: cell-e29a42fbdc3ccdc5, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8 [Note: Includes Hidden Tests]

assert calc_a
assert type(calc_a) == float

### BEGIN HIDDEN TESTS
assert calc_a == 11.0
### END HIDDEN TESTS
```

<!-- grade_id: cell-00d84e7f456e3984, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8 [Note: Includes Hidden Tests]

assert calc_b
assert type(calc_b) == float

### BEGIN HIDDEN TESTS
assert calc_b == 2.0
### END HIDDEN TESTS
```

<!-- grade_id: cell-c5231a0cfa605a94, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8 [Note: Includes Hidden Tests]

assert calc_c
assert type(calc_c) == float

### BEGIN HIDDEN TESTS
assert calc_c == 78.0
### END HIDDEN TESTS
```

<!-- grade_id: cell-9b225dabad981fb1, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q8 [Note: Includes Hidden Tests]

assert calc_d
assert type(calc_d) == float

### BEGIN HIDDEN TESTS
assert calc_d == 2.5
### END HIDDEN TESTS
```

### Q9 - Curve Grade  (1 point)

Write a function to apply a very simple curve to an input grade. The function will be called `curve_grade` and will take two parameters as inputs, in the order specified here: 

- `points` : the number of points a student earned on an assignment (int or float)
- `curve` : the proportion of curve to apply to the student's `points` (int or float)

Within the function, the code will `return` `curved_points`. This code will apply the proportion specified in `curve` to the student's points. For example, if `curve` is 0.1 and `points` is 100, the code in the funtion would multiply `100` (`points`) by `0.1` (`curve`) and add that back to the original `points`, `return`ing the value `110`. If `points` were `90` and `curve` was `0.05`, the function would `return` `94.5`.

<!-- grade_id: cell-1328d643463d0a68, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def curve_grade(points, curve):
    
    curved_points = points + (points * curve)
    
    return curved_points
### END SOLUTION
```

```python
curve_grade(10, 0.1)
```

Output:
```
11.0
```

```python
# BE SURE YOUR ANSWER IS IN THE CELL ABOVE
# but you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-ff0f16aa91a10950, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9

assert curve_grade
assert curve_grade(100, 0.1) == 110
```

<!-- grade_id: cell-6a372d4477ec3403, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q9 [Note: Hidden Tests]

### BEGIN HIDDEN TESTS
# tests parameter names are correct
assert curve_grade(points=100, curve=0.1) == 110
### END HIDDEN TESTS
```

<!-- grade_id: cell-e74d5a799a1776cc, points: 0.5, grade: true, solution: false -->
```python
## TESTS FOR Q9 [Note: Hidden Tests]

### BEGIN HIDDEN TESTS
# tests other input values
assert curve_grade(points=100, curve=0.05) == 105
assert curve_grade(points=90, curve=0.05) == 94.5
assert curve_grade(points=90, curve=0.1) == 99.0
### END HIDDEN TESTS
```

### Q10 - Using `curve_grade`  (0.75 points)

Execute your `curve_grade` function to generate the `curved_points` for 3 students in the same class on their first exam, applying a 2% curve (0.02 for `curve`) to each student's score. Store the student's curved grade in the variable specified below:

- `student_a` | earned 12.2 `points`
- `student_b` | earned 11 `points`
- `student_c` | earned 11.5 `points`

<!-- grade_id: cell-dd63cd542096825b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
student_a = curve_grade(12.2, 0.02)
student_b = curve_grade(11, 0.02)
student_c = curve_grade(11.5, 0.02)
print(student_a, student_b, student_c)
### END SOLUTION
```

Output:
```
12.443999999999999 11.22 11.73
```

<!-- grade_id: cell-27e516bea517e29f, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q10 [Note: Hidden Tests]

assert student_a
assert type(student_a) == float

### BEGIN HIDDEN TESTS
assert math.isclose(student_a, 12.44, rel_tol = 0.01)
### END HIDDEN TESTS
```

<!-- grade_id: cell-33337142852a443b, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q10 [Note: Hidden Tests]

assert student_b
assert type(student_b) == float

### BEGIN HIDDEN TESTS
assert math.isclose(student_b, 11.22, rel_tol = 0.01)
### END HIDDEN TESTS
```

<!-- grade_id: cell-76452dd79b41edc4, points: 0.25, grade: true, solution: false -->
```python
## TESTS FOR Q10 [Note: Hidden Tests]

assert student_c
assert type(student_c) == float

### BEGIN HIDDEN TESTS
assert math.isclose(student_c, 11.73, rel_tol = 0.01)
### END HIDDEN TESTS
```

## The End!

This is the end of the assignment!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. When you are ready, submit on datahub!