# A2-Ciphers

This assignment is out of 8 points, worth 7% of your grade.

This assignment covers Conditionals, Collections, and Functions. It builds up to the application of Ciphers. 

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.**

## Overview

Make sure you delete `raise NotImplementedError()` whenever you see it in the assignment, replacing it with your code for the question. To have the best chance at receiving full credit, be sure your completed assignment passes all the asserts silently. This assignment also has hidden tests - which means that passing all the asserts you see does not guarantee that you have the correct answer! Make sure to double check and re-run your code to make sure it does what you expect, and that the variables you define end up with values you'd expect. 

## Part 1 - Conditionals

### Q1 - Conditionals With Booleans (0.5 points)

Write a function `check_status()` that uses a conditional that takes a single boolean parameter `status` as input. Inside the function:
- if `status` evaluates as `True`, set the variable `output` as the string 'GOOD'
- else (meaning `status` must evaluate as `False`), set the variable `output` as the string 'BAD'

Be sure to `return` `output` from your function.

<!-- grade_id: cell-d2e812dc1f9a43a6, points: 0, grade: false, solution: true -->
```python
def check_status(status):
    if status:
        output = 'GOOD'
    else:
        output = 'BAD'
    
    return output

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-de368a5bcd2db9c1, points: 0.25, grade: true, solution: false -->
```python
assert callable(check_status)
assert type(check_status(status=True)) == str
assert check_status(status=True) == 'GOOD'
```

<!-- grade_id: cell-60a753f009ed6b1b, points: 0.25, grade: true, solution: false -->
```python
assert callable(check_status)
assert type(check_status(False)) == str
assert check_status(False) == 'BAD'
```

### Q2 - Conditionals With Comparisons (0.5 points)

Consider a case in which you are running an experiment with multiple groups of subjects, including young subjects and old subjects. 

Based on the subject's age group, you want to do a different analysis if they are in the young group,.

Write a function `change_analysis()`  that takes a single parameter `subj_age`. Inside the function:
- Check `if` should check if the variable `subj_age` is equal to the value `'young'` (a string).
    - If so, it should set the variable `change` to the value `True` (a boolean)
- `else`, the variable `change` should be set to `False` (a boolean)

Be sure to `return` the variable `change` from the function

<!-- grade_id: cell-5fdac280e980848a, points: 0, grade: false, solution: true -->
```python
def change_analysis(subj_age):
    if subj_age == 'young':
        change = True
    else:
        change = False
    
    return change

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-2e59c9f37e0310a5, points: 0.25, grade: true, solution: false -->
```python
assert callable(change_analysis)
assert change_analysis(subj_age='young') == True
```

<!-- grade_id: cell-c4621a021d61b90f, points: 0.25, grade: true, solution: false -->
```python
assert callable(change_analysis)
assert change_analysis(subj_age='old') == False
```

### Q3 - Conditionals With Multiple Options (1 point)

Sometimes we may need to check multiple conditions, with many (more than two) possibilities.

Define a function `check_conditions()` that takes two parameters as input: `check_1` and `check_2`.

Within the function, carry out the following operation:
- if `check_1` and `check_2` variables are both True, it should set the value of a variable `outcome` to the string 'BOTH'
- else if `check_1` is True and `check_2` is False, it should set the value of a variable `outcome` to the string 'ONE'
- else if `check_1` is False and `check_2` is True, it should set the value of a variable `outcome` to the string 'TWO'
- else (meaning both must be False), it should set the value of a variable `outcome` to the string 'NEITHER'

Be sure to `return` `outcome` from the function.

<!-- grade_id: cell-0cada564ef3bd740, points: 0, grade: false, solution: true -->
```python
def check_conditions(check_1, check_2):
    if check_1 and check_2:
        outcome = 'BOTH'
    elif check_1 and not check_2:
        outcome = 'ONE'
    elif not check_1 and check_2:
        outcome = 'TWO'
    else:
        outcome = 'NEITHER'
    
    return outcome

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-7b81b816f5fbd8fc, points: 0.25, grade: true, solution: false -->
```python
# Check 'BOTH' case
assert callable(check_conditions)
assert type(check_conditions(True, True)) == str
assert check_conditions(check_1=True, check_2=True) == 'BOTH'
```

<!-- grade_id: cell-621a9631f8689eb8, points: 0.25, grade: true, solution: false -->
```python
# Check 'ONE' case
assert check_conditions(True, False) == 'ONE'
```

<!-- grade_id: cell-a01f95735ce601a7, points: 0.25, grade: true, solution: false -->
```python
# Check 'TWO' case
assert check_conditions(False, True) == 'TWO'
```

<!-- grade_id: cell-74ded29c5e0ab108, points: 0.25, grade: true, solution: false -->
```python
# Check 'NEITHER' case (using hidden test)

```

### Q4 - Multiple Cases: Experiment Example (1 point)

Let's revisit the setup of Q2, in which we wanted to set a variable to indicate if a subject in an experiment was in a particular group, and thus required different processing. In this question we will extend that premise by writing a set of conditionals to mark an interaction between subject's age and score. 

For example, it is common for behavioral performance to decrease with age. However, some subjects do not seem to show this behavioral deficit with age. Investigating such 'super-agers' could be useful for learning about the cognitive decline of aging.

Write a function `determine_ager()` that has two parameters `subj_age` and `score`.

Inside the function utilize a conditional that acomplishes the following:
- if the variable `subj_age` has the value `old` and the variable `score` is greater than 25:
    - set the variable `subj_status` to the string 'Super Ager'
    
    
- else if the variable `subj_age` has the value `old` and the variable `score` is less than or equal to 25:
    - set the variable `subj_status` to the string 'Normal Ager'
    
    
- else if the variable `subj_age` has the value `young` and the variable `score` is greater than 25:
    - set the variable `subj_status` to the string 'Normal Youth'
    
    
- else if the variable `subj_age` has the value `young` and the variable `score` is less than or equal to 25:
    - set the variable `subj_status` to the string 'Bad Youth'
    
Be sure to `return` `subj_status` from the function

<!-- grade_id: cell-514a6e20b819dee0, points: 0, grade: false, solution: true -->
```python
def determine_ager(subj_age, score):
    if subj_age == 'old' and score > 25:
        subj_status = 'Super Ager'
    elif subj_age == 'old' and score <= 25:
        subj_status = 'Normal Ager'
    elif subj_age == 'young' and score > 25:
        subj_status = 'Normal Youth'
    elif subj_age == 'young' and score <= 25:
        subj_status = 'Bad Youth'
    
    return subj_status

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-10e577dabe2f9a4f, points: 0.25, grade: true, solution: false -->
```python
subj_age = 'old'
score = 28
ex1 = determine_ager(subj_age, score)

assert callable(determine_ager)
assert type(ex1) == str
assert ex1 == 'Super Ager'
```

<!-- grade_id: cell-c259aac343bb297a, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b05b4bec57b61e00, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0cc09b0683fdfd57, points: 0.25, grade: true, solution: false -->
```python

```

### Q5 - Creating a Calculator (1.25 points)

Finally (for this section), we will write a flexible calculator that performs multiple operations (addition, subtraction, multiplication, and division) on pairs of numbers. 

To do so, define a function `calculate()` with a three parameters:
- `operation` | The value specified for the `operation` parameter on execution will indicate which operation to perform
- `num1` | The first value to be used in the calculation
- `num2` | The second value to be used in the calculation

The calculator will assign the output of its calculation to a variable called `answer`, which will be `return`ed from the function. 

Specifically, create the `calculate()` function such that:
- if `operation` has the value 'add', then `answer` gets set as the sum of `num1` and `num2`
- elif `operation` has the value 'subtract', then `answer` gets set as the difference of `num2` from `num1`
- elif `operation` has the value 'multiply', then `answer` gets set as the product of `num1` and `num2`
- elif `operation` has the value 'divide', then `answer` gets set as the quotient of `num1` divided by `num2`
- else `answer` gets set as the value 'not a thing'

For an example, if we were to execute our function as follows:

```python
output = calculate(operation='multiply', num1=2, num2=3)
```

we would expect `output` to store `6`.

<!-- grade_id: cell-d44d4aa343270410, points: 0, grade: false, solution: true -->
```python
def calculate(operation, num1, num2):
    if operation == 'add':
        answer = num1 + num2
    elif operation == 'subtract':
        answer = num1 - num2
    elif operation == 'multiply':
        answer = num1 * num2
    elif operation == 'divide':
        answer = num1 / num2
    else:
        answer = 'not a thing'
    
    return answer

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-63991f35fb0fa2b7, points: 0.25, grade: true, solution: false -->
```python
# test addition
assert callable(calculate)
assert type(calculate(operation='add', num1=2, num2=3)) == int
assert calculate(operation='add', num1=2, num2=3) == 5
assert calculate(operation='add', num1=10, num2=-5) == 5
```

<!-- grade_id: cell-df854472375eaaea, points: 0.25, grade: true, solution: false -->
```python
# test subtraction
assert calculate(operation='subtract', num1=2, num2=3) == -1
assert calculate(operation='subtract', num1=5, num2=5) == 0

```

<!-- grade_id: cell-5cb26332c714b7b6, points: 0.25, grade: true, solution: false -->
```python
## Hidden tests to check multiplication operation is correct

```

<!-- grade_id: cell-515fa93f1045b72f, points: 0.25, grade: true, solution: false -->
```python
## Hidden tests to check division operation is correct

```

<!-- grade_id: cell-544aefbe440ec548, points: 0.25, grade: true, solution: false -->
```python
## Hidden test to check unrecognized operation

```

## Part 2: Collection Types

This part covers defining and working with the collection types: lists, tuples, and dictionaries. 

### Q6 - Lists (0.25 points)

Below are the results of a hypothetical experiment of measuring the height of a class in a land far away:

- Vanessa is 5.6 feet tall
- Sarah is 5.4 feet tall
- Demi is 6.1 feet tall
- George is 5.9 feet tall
- Foram is 5.5 feet tall

Create a list called `class_names` and fill it with the names of each person in the class (as strings). 

Create another list called `class_heights` and fill it with the heights of each person in the class (as floats). 

Remember that the order of lists and the spelling of the names matter; be sure to keep the order of items the same as listed above and the spelling the same!

<!-- grade_id: cell-653ac226fd7e65c4, points: 0, grade: false, solution: true -->
```python
class_names = ['Vanessa', 'Sarah', 'Demi', 'George', 'Foram']
class_heights = [5.6, 5.4, 6.1, 5.9, 5.5]

```

<!-- grade_id: cell-7ed89774a93617f3, points: 0.25, grade: true, solution: false -->
```python
# Check both lists are defined
assert type(class_names) == list
assert type(class_heights) == list

# Check both lists have the same length, and are of length 5
assert len(class_names) == len(class_heights) == 5


```

Now that we have the data stored in lists, we will explore other collection types. 

### Q7 - Tuples (0.25 points)

Create a tuple that contains just the names of the people in the experiment. Call this tuple `names_tuple`. 

<!-- grade_id: cell-fb210f0469973e5a, points: 0, grade: false, solution: true -->
```python
names_tuple = ('Vanessa', 'Sarah', 'Demi', 'George', 'Foram')

```

<!-- grade_id: cell-5969238aefebc1dd, points: 0.25, grade: true, solution: false -->
```python
assert type(names_tuple) == tuple


```

### Q8 - Dictionaries (0.25 points)

Next, store the data in a dictionary where each student's name is the key and their height is the value.

Call this dictionary `results_dictionary`. 

<!-- grade_id: cell-d1afd11046ada78c, points: 0, grade: false, solution: true -->
```python

results_dictionary = {'Vanessa': 5.6, 'Sarah': 5.4, 'Demi': 6.1, 'George': 5.9, 'Foram': 5.7}

```

<!-- grade_id: cell-a823da59fce8614b, points: 0.25, grade: true, solution: false -->
```python
assert type(results_dictionary) == dict


```

### Q9 - Lists of Lists (0.5 points)

You can also make lists that are filled with lists! List-ception. 

First, create three different lists:
- A list called `string_list` that contains three strings (can be any strings)
- A list called `number_list` that contains three numbers (can be any numbers - int or float)
- A list called `boolean_list` that contains three boolean (can be any booleans)

Then, create a new list, called `nested_list` which contains the three lists you created above (in the order specified in the above bullet points). 

<!-- grade_id: cell-49a40611dcdbbea0, points: 0, grade: false, solution: true -->
```python

string_list = ['a', 'b', 'c']
number_list = [12.4, 16.4, 11.2]
boolean_list = [True, False, True]
nested_list = [string_list, number_list, boolean_list]

```

<!-- grade_id: cell-51ffc8b55e7f36fc, points: 0.25, grade: true, solution: false -->
```python
assert type(string_list) == list
assert len(string_list) == 3
assert type(string_list[0]) == str

assert type(number_list) == list
assert len(number_list) == 3
assert type(number_list[0]) == int or type(number_list[0]) == float

assert type(boolean_list) == list
assert len(boolean_list) == 3
assert type(boolean_list[0]) == bool
```

<!-- grade_id: cell-04845071d9603b3d, points: 0.25, grade: true, solution: false -->
```python
assert type(nested_list) == list
assert len(nested_list) == 3
assert type(nested_list[0]) == list
```

### Q10 - Indexing (0.5 points)

Do the following *using indexing*:
- Store the first value in `list_1` to `index_1`
- Store the third value from the end of `list_2` to `index_2`
- Store the last three values of `list_4` to `index_3` 
- Store the first three values of `list_3` to `index_4`

Notes: 
- Values must be stored in the same order they appear in the original list in your output.
- Be sure to execute the cell below to define the lists prior to referencing them in your answer.

```python
# For Q10 & Q11, the following lists are provided for you
list_1 = [10, 20, 30, 40]
list_2 = [13, 15, 17, 19, 21, 23]
list_3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
list_4 = [21, 9, 98, 289, 938]
```

<!-- grade_id: cell-650e61f820451347, points: 0, grade: false, solution: true -->
```python

index_1 = list_1[0]
index_2 = list_2[-3]
index_3 = list_4[-3:]
index_4 = list_3[0:3]

```

<!-- grade_id: cell-2468c24a17df0b96, points: 0.125, grade: true, solution: false -->
```python
assert type(index_1) == int


```

<!-- grade_id: cell-d2488aa0f0a84e65, points: 0.125, grade: true, solution: false -->
```python
assert type(index_2) == int


```

<!-- grade_id: cell-3ae0d083dd79537d, points: 0.125, grade: true, solution: false -->
```python
assert type(index_3) == list
assert len(index_3) == 3


```

<!-- grade_id: cell-f591dd9f036b0ffe, points: 0.125, grade: true, solution: false -->
```python
assert type(index_4) == list
assert len(index_4) == 3


```

### Q11 - Comparisons using Indexing (0.5 points)

Do the following comparisons *using indexing*, storing the output in the variable at left:
- `comp_result_1` | Determine if the third value of `list_1` is one of the last three values in `list_2`
- `comp_result_2` | Determine if the second value of `list_4` is greater than the third value in `list_3`
- `comp_result_3` | Determine if the last value of `list_3` is less than the second value in `list_2`
- `comp_result_4` | Determine if the second value of `list_1` multiplied by the fifth value of `list_3` is less than the penultimate value in `list_2`
    
Keep in mind that you are storing the result of a comparison to a variable, so resulting variables should all be booleans. 

<!-- grade_id: cell-58d83066fbe5dfa8, points: 0, grade: false, solution: true -->
```python
comp_result_1 = list_1[2] in list_2[-3:]
comp_result_2 = list_4[1] > list_3[2]
comp_result_3 = list_3[-1] < list_2[1]
comp_result_4 = (list_1[1]*list_3[4]) < list_2[-2]

```

<!-- grade_id: cell-2209a4edc72a8364, points: 0.125, grade: true, solution: false -->
```python
assert type(comp_result_1) == bool


```

<!-- grade_id: cell-6353d9279a80e98f, points: 0.125, grade: true, solution: false -->
```python
assert type(comp_result_2) == bool


```

<!-- grade_id: cell-d75011cb18874b17, points: 0.125, grade: true, solution: false -->
```python
assert type(comp_result_3) == bool


```

<!-- grade_id: cell-244b969301ac73e8, points: 0.125, grade: true, solution: false -->
```python
assert type(comp_result_4) == bool


```

## Part 3:  Ciphers

This part is an application of all the tools we have covered so far, to the problem of creating and using ciphers. 

A cipher is a procedure for encoding and decoding secret messages (or encrypted data). 

To work on this application, we are going to need to use functions and operators. 

It also uses concepts and ideas from encoding and decoding - and in particular the operators `ord` and `chr`. 

```python
# Example / outline of writing a cipher
#   The following shows a layout of writing a cipher.
#   You will be writing code to actually do the encoding / decoding.

# Create a message
original_message = 'Hello from Prof'
print('Original Message:\t', original_message)

# First, Apply Encoding Code, to convert to an encoded message
encoded = 'ĐĭĴĴķèĮĺķĵèĜķĵöĮ'
print('Encoded Message:\t', encoded)

# Then, Apply Decoding Code, to decode into an interpretable message
decoded = 'Hello from Prof'
print('Decoded Message:\t', decoded)
```

Output:
```
Original Message:	 Hello from Prof
Encoded Message:	 ĐĭĴĴķèĮĺķĵèĜķĵöĮ
Decoded Message:	 Hello from Prof
```

### Q12 - Creating an Encoder (0.25 points)

Write a function `encoder()` to encode a character from its input character into its secret message character.

Your function will take two parameters:
- `char` | The input character to be encoded (str)
- `key` | The key to be used during encoding (int; default: 200)

Within the function, convert `char` to another character `encoded_char`, using the following:
- Get the unicode code point for the `char` (using `ord`)
- Add the value of `key` to that code point (which should be an int)
- Convert this new int back to a character (using `chr`).

Be sure to `return` `encoded_char` from the function

<!-- grade_id: cell-d77bc7ed64441ebc, points: 0, grade: false, solution: true -->
```python
def encoder(char, key=200):
    
    encoded_char = chr(ord(char) + key)
    
    return encoded_char

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-8dbc526034a66ab8, points: 0.1, grade: true, solution: false -->
```python
# test encoder with default value for key
assert encoder('h') == 'İ'
assert encoder('e') == 'ĭ'
assert encoder('l') == 'Ĵ'
assert encoder('l') == 'Ĵ'
assert encoder('o') == 'ķ'
```

<!-- grade_id: cell-fb95b9e70589f512, points: 0.15, grade: true, solution: false -->
```python
# hidden test to check for functionality with other values for key

```

### Q13 - Decoder (0.25 points)

Write a function `decoder()` to decode a secretly-encoded character back to its original character.

Your function will take two parameters:
- `encoded_char` | The enoded character to be decoded (str)
- `key` | The key to be used during encoding (int; default: 200)

Within the function, convert `encoded_char` back to `char` by:
- Getting the unicode code point for `encoded_char` (using `ord`)
- Subtracting the value of `key` from that code point (which should be an int)
    - This undoes the secret encoding of the character, getting back to the original value
- Convert this new int back to a character (using `chr`).

The decoded character should be stored in `char`, which should be `return`ed from the function.

Note that the code to answer this question should look very similar to how you answered the encoding question. In fact, you can even copy over the encoding code, and make some small changes to make it function as a decoder. 

<!-- grade_id: cell-b629dac4beccd5f1, points: 0, grade: false, solution: true -->
```python
def decoder(encoded_char, key=200):

    char = chr(ord(encoded_char) - key)
    
    return char

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-ddf1c25663771ad8, points: 0.1, grade: true, solution: false -->
```python
# tests decoder with default key value
assert decoder('İ') == 'h'
assert decoder('ĭ') == 'e'
assert decoder('Ĵ') == 'l'
assert decoder('Ĵ') == 'l'
assert decoder('ķ') == 'o'
```

<!-- grade_id: cell-25981e96494c9450, points: 0.15, grade: true, solution: false -->
```python
# hidden test to check for functionality with other values for key

```

### Using a Custom Encoder

The encoder & decoder we created above is built on conversions based on manipulating unicode code points.

If a potential attacker wanted to decode our messages, and happened to know we were using this approach, there are some relatively simple ways that they could try to figure out our encoding key, and then, if successful, decode our messages. 

Thought experiment: can you think of ways to decode messages that use this encoding?

To thwart this potential attacker, let's now explore using a *custom encoding*. Instead of manipulating code points, we are going to use a custom mapping to switch out characters to encode our messages. 

### Q14 - Custom Encoder (1 point)

Write a function `custom_encoder()` to encode a character from original messages, using our custom encoding scheme (defined below).

This function should take two input parameters:

- `char` | the character to be encoded using our custom encoding scheme
- `key` | the `key` to be used in encoding (int; default: 200)

Within the function, carry out the following:
- Define the dictionary `custom_encodings` (provided at end of this cell)
- If `char` is a key in the `custom_encodings` dictionary, store the corresponding value from the dictionary in `output_char` (i.e. if `char` were 'o', the function would return `'n'`)
- Otherwise, use the `encoder()` function you wrote earlier, storing the output in `output_char`, being sure to pass both `char` and `key` in as inputs.

Be sure to `return` the encoded character (`output_char`) from the function
        
        
```python
custom_encodings = {
    'a' : 'r', 
    'e' : 'p', 
    'i' : 'm',
    'o' : 'n',
    'u' : 's'
}
```

<!-- grade_id: cell-c71c52f6393468cd, points: 0, grade: false, solution: true -->
```python
def custom_encoder(char, key=200):

    custom_encodings = {
        'a' : 'r', 
        'e' : 'p', 
        'i' : 'm',
        'o' : 'n',
        'u' : 's'
    }
    
    if char in custom_encodings:
        output_char = custom_encodings[char]
    else:
        output_char = encoder(char=char, key=key)
        
    return output_char

```

<!-- grade_id: cell-66b937622cc2a5f6, points: 0.5, grade: true, solution: false -->
```python
# test for char in custom_encodings
custom_encoded = custom_encoder('a')
assert type(custom_encoded) == str
assert custom_encoded == 'r'


```

<!-- grade_id: cell-36ff28ecfc316fc0, points: 0.5, grade: true, solution: false -->
```python
# tests for char NOT in custom_encodings & variable key
custom_encoded = custom_encoder('A')
assert type(custom_encoded) == str
assert custom_encoded == 'ĉ'


```

## The End!

This is the end of the assignment!

Have a look back over your answers, and also make sure to Restart & Run All from the kernel menu to double check that everything is working properly and all of your asserts pass silently. 

When you are ready, submit to datahub.