# Assignment 3: Chatbots

This assignment is out of 8 points, worth 7% of your grade.

This assignment focuses on loops and functions and builds up to the application of Chatbots. 

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.**

## Overview

Make sure you delete `raise NotImplementedError()` whenever you see it in the assignment, replacing it with your code for the question. To have the best chance at receiving full credit, be sure your completed assignment passes all the asserts. This assignment also has hidden tests - which means that passing all the asserts you see does not guarantee that you have the correct answer! Make sure to double check and re-run your code to make sure it does what you expect.

### Quick Tips

Make sure you execute your code to define the functions you write, and then add a new cell to apply your function to different inputs, and check that you get the outputs you expect. 

If you're not sure where to start, look at the code in the test cases (in the asserts) for examples of what kind of inputs and outputs the function you are writing should take, and what kind of behavior it should have.

Keep in mind that each question is mostly independent from every other - so if you get stuck on something, you can always move on to the next one, and come back to any that you have trouble with later.

## Setup

#### Code Imports

<span style="color:red">You must run the following cell before proceeding with the assignment.</span>

```python
# Setup - run this cell before doing the next part of the assignment
#   This imports some extra code, making it available for us to use later
import string
import random
import nltk
```

### Zen Chatbot

In this assignment, we will be working our way up to building a chatbot.

To start with, we can have a quick chat with a zen chatbot to see what these kinds of chatbots look like. 

Also, if you get stuck on anything, come back here and talk to the Zen chatbot to relax :)

Note: in this, and other chatbot lines, running this code will start an open-ended while loop, that will interrupt your code flow until you explicitly quit out of the conversation. If you want to be able to run all your code (from kernel `restart & run all`), then comment these lines back out when you are finished exploring them. 

```python
# Uncomment and run the line of code in this cell to talk to a Zan Chatbot
# when you're done with your chat, type 'quit'
# nltk.chat.zen.demo()
```

## Part 1: Loops

This part covers using loops. 

Note that you should NOT define functions for the questions in this section.

### Q1 - `for` Loop (0.5 points)

What is the sum of all even numbers between 20 and 50 (inclusive)?

Answer this question using a for loop, and using `range`. 

Save the result out to a variable called `sum_result`. 

<!-- grade_id: cell-749ae260cbdc082d, points: 0, grade: false, solution: true -->
```python
sum_result = 0
for number in range(20, 51, 2):
    sum_result = sum_result + number

```

<!-- grade_id: cell-6b415de65a9332be, points: 0.5, grade: true, solution: false -->
```python
assert type(sum_result) == int


```

### Q2 - `while` Loop (0.5 points)

Write a while loop that adds the first 50 numbers (from 1 to 50, inclusive) together.

To do so, use a `while` loop that continues while the condition of `cur_num` being less than or equal to `stop_num` is `True`. 

Inside the loop, it should add `cur_num` to the running `total` variable, and then add 1 to `cur_num`. 

Make sure to check at the end that the value of `total` reflects the total summation. 

<!-- grade_id: cell-fdabc8d0323473dd, points: 0, grade: false, solution: true -->
```python
# These variables provided for you - use them in your solution
total = 0
cur_num = 1
stop_num = 50

while cur_num <= stop_num:
    total = total + cur_num
    cur_num = cur_num + 1

```

<!-- grade_id: cell-f6513546d3afb423, points: 0.25, grade: true, solution: false -->
```python
assert type(cur_num) == int
assert type(total) == int

# This will check that the `cur_num` increment variable ended at the right value
assert cur_num == 51
```

<!-- grade_id: cell-16673b5754d8321e, points: 0.25, grade: true, solution: false -->
```python
# hidden test to check the corret value is stored  in total

```

### Example: `append`

The following cell gives an example of the `append` method for lists that we discussed in lecture as a reminder:

To add an item, `my_item`, to the end of an existing list, `my_list`, you can use `my_list.append(my_item)`

```python
# Example of appending to a list
my_list = ['a', 12, None]

print('Before Append: ', my_list)

# append True to my_list
my_list.append(True)

print('After  Append: ', my_list)
```

```python
# Example of using append with loops
list_of_values = [0, 1, 2, 3]

# Initialize an empty list that will be used to store outputs
outputs = []

# Loop through all items in `list_of_items`
for item in list_of_values:
    
    # Add incremented item to the list of outputs
    outputs.append(item + 1)

# Check what outputs ends up as
print(outputs)
```

### Q3 - Loops: Multiplication (0.5 points)

Using the provided variable `num_list`, write a for loop to loop across each value, multiplying it by -1. 

Store the result for each value in a list called `eval_1` by defining `eval_1` before the loop and using `append` inside the loop. 

```python
# This creates a list of all numbers from 1-12, inclusive
num_list = list(range(1, 13))
```

<!-- grade_id: cell-c3c8595879347f24, points: 0, grade: false, solution: true -->
```python
eval_1 = []
for item in num_list:
    eval_1.append(item * -1)

```

<!-- grade_id: cell-09e1d8d99d175d45, points: 0.5, grade: true, solution: false -->
```python
assert type(eval_1) == list
assert type(eval_1[0]) == int


```

### Example: Looping through Dictionaries

```python
# This example of looping through dictionaries provided
dictionary = {
    'key1' : 'value1', 
    'key2' : 'value2'
}

for item in dictionary:
    
    # When you loop through a dictionary like this, `item` is the key 
    key = item
    
    # Using the key, you can access the associated value of the dictionary from within the loop
    value = dictionary[item]
    
    # Print out data for each loop iteration
    print('New Loop Iteration')
    print('\tKey is:\t', item)
    print('\tVal is:\t', value)
```

### Q4 - Loops: Dictionaries (0.5 points)

Using the `subjects` dictionary, write a for loop that loops across the dictionary and collects all subject numbers (ex. 'S1') where the dictionary value is True. 

Imagine, for example, the dictionary indicates whether processing is complete, and we wanted to get a list of the subjects where processing is complete. 

To answer this question, use a for loop across the `subjects` dictionary. You then need to get the associated value in each iteration and check if it is `True`. If it is `False`, you can use `continue` to skip ahead to the next iteration. Otherwise, append the subject number (ex 'S2') to a list called `finished`. 

```python
# This dictionary provided for you
subjects = {
    'S1' : True,
    'S2' : False,
    'S3' : True,
    'S4' : False,
}
```

<!-- grade_id: cell-1118d6a4ac43e87d, points: 0, grade: false, solution: true -->
```python
finished = []
for subject in subjects:
    
    if subjects[subject] == False:
        continue
    else:
        finished.append(subject)

```

<!-- grade_id: cell-36d719f579df1fe3, points: 0.25, grade: true, solution: false -->
```python
assert type(finished) == list
assert len(finished) == 2
```

<!-- grade_id: cell-808d35d9c1be0997, points: 0.25, grade: true, solution: false -->
```python
# hidden test to check that the correct list is returned

```

## Part 2: Functions

Here, we'll define and execute a number of functions, a number of which will be used in Part 3

### Q5 - Functions & Loops: Division (0.5 points)

Write a function `div_3()` that will find all the values from the input parameter `data_list` that are perfectly divisible by 3. 

Use a for loop with a conditional to do this inside your function, and save any values from `data_list` that meet this condition into a new list `output`, which is `return`ed from the function.  

<!-- grade_id: cell-b3adf3ba9e29eb9f, points: 0, grade: false, solution: true -->
```python
def div_3(data_list):
    output = []
    
    for number in data_list:
        if number % 3 == 0:
            output.append(number)
            
    return output


```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-bb9b4796b7db8253, points: 0.5, grade: true, solution: false -->
```python
# define list to use in testing
test_list = [-1, 20, -100, -44, 33, 97, -101, 45, -79, 96]

assert callable(div_3)
assert type(div_3(test_list)) == list
assert len(div_3(test_list)) == 3


```

### Q6 - Functions & Loops: Odd or Even (0.5 points)

Write a function `is_even()` which will check whether the values in your input parameter (`data_list`) are even or odd. 

Specifically, for each value, it should add `True` to a new list `output_list` if the value is even, and `False` to `output_list` if the value is odd. 

Be sure to `return` `output_list` from the function.

<!-- grade_id: cell-af4c9eabb79396f1, points: 0, grade: false, solution: true -->
```python
def is_even(data_list):
    output_list = []
    for item in data_list:
        if item % 2 == 0:
            output_list.append(True)
        else:
            output_list.append(False)
    
    return output_list

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-3b8896dd294ccc5a, points: 0.5, grade: true, solution: false -->
```python
# define list to use in testing
test_list = [-1, 20, -100, -44, 33, 97, -101, 45, -79, 96]

assert callable(is_even)
assert len(is_even(test_list)) == len(test_list)


```

### Q7 - Find Max Value  (0.5 points)

Write a function called `find_max` that takes in a single parameter called `random_list`, which should be a list. 

This function will find the maximum (max) value of the input list, and `return` it. 

This function will assume the list is composed of only positive numbers.

To find the max, within the function, use a variable called `list_max` that you initialize to value 0. 

Then use a `for` loop to loop through `random_list`. 

Inside the loop, use a conditional to check if the current value is larger than `list_max`, and if so, re-assign `list_max` to store this new value.

After the loop, return `list_max`, which should now store the maximum value from within the input list. 

<!-- grade_id: cell-23d051fdc803f0bd, points: 0, grade: false, solution: true -->
```python
def find_max(random_list):
    
    list_max = 0

    for number in random_list:
        
        if number > list_max:
            list_max = number
            
    return list_max

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-404d30e1e65b6b83, points: 0.25, grade: true, solution: false -->
```python
assert callable(find_max)
assert type(find_max([1, 7, 5])) == int
assert find_max([1, 7, 5]) == 7
```

<!-- grade_id: cell-f39da4732d3c6281, points: 0.25, grade: true, solution: false -->
```python
assert callable(find_max)


```

### Q8 - Square Everything  (0.5 points)

Write a function called `square_all` that takes an input called `collection` that you assume to be a list or tuple of numbers. 

This function should return a new list, which contains the square of each value in `collection`. 

To do so, inside the function, you will define a new list, use a loop to loop through the input list, use an operator to square the current value, and use the `append` method to add this to your output list. After the loop, return the output list.

<!-- grade_id: cell-5993fedb7f503a6f, points: 0, grade: false, solution: true -->
```python
def square_all(collection):
    
    square_list = []
    for val in collection:
        square_list.append(val**2)
    
    return square_list

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-8a86a6f3ec803f9c, points: 0.25, grade: true, solution: false -->
```python
assert callable(square_all)
assert type(square_all([1, 2])) == list
assert square_all((2, 4)) == [4, 16]
```

<!-- grade_id: cell-09d44aca37e9d05f, points: 0.25, grade: true, solution: false -->
```python
assert callable(square_all)


```

## Part 3:  Chatbots

This part is an application of all the tools we have covered so far, to the problem of creating a chatbot.

A chatbot is a piece of software that tries to have a conversation with a human.

While today's chatbots are powerful learners, traditional chatbots are much simpler. 

Traditional chatbots follow a "rule-based" approach that is defined for them (by you, the programmer), usually without any learning. 

These rules can be very simple or sometimes quite complex - but everything is pre-defined. 

These types of chatbots can handle simple conversations but do tend to fail at more complex ones.

#### Chatbot functions 

In order to make our chatbot, we will create a series of functions that our eventual chatbot will use together to try and chat with a user. 

At the end, we will test that these functions do indeed allow us to have simple conversations with our chatbot.

Below are some more example chatbots to talk to. Remember to comment these lines back out if you want to be able to `Restart and Run All`.

### Example Chatbot - Eliza

The following demo is from the Natural Language ToolKit (nltk), a Python package available through Anaconda. 

Uncomment the line below, and run the code to be able to have a chat with Eliza.

More information on Eliza: https://en.wikipedia.org/wiki/ELIZA

```python
# Run a demo of a chatbot, with Eliza
# nltk.chat.eliza.demo()
```

### Example Chatbot - Iesha

The following is another demo chatbot from nltk, this one called Iesha. 

```python
# Run a demo of a chatbot, with Iesha
# nltk.chat.iesha.demo()
```

As you try the chatbot demos with Eliza and Iesha, see if you can notice any patterns that might help you figure out how they work. 

That is, can you guess at what kind of programming structures these programs might use?

### Chatbots: Words & Rules

So, how do these kinds of chatbots work? Most broadly, they combine two simple but powerful approaches:
- Rule Based Artificial Intelligence (GOFAI)
- Bag-Of-Word Language Models

**Rule based AI**, sometimes called 'Good Old-Fashioned Artificial Intelligence', or 'GOFAI', is an approach to making artificial agents that tries to write down a set of rules to follow to perform some task. In practice, this often looks like using a lot of conditionals: if `X` happens, do `Y`. If we can write down enough logical procedures (call them algorithms) to follow to complete some task, then we can create an artificially intelligent agent (for that task).

**Bag-Of-Word language models** are an approach to language that takes the simple approach of saying 'let's try and figure out what someone is saying simply by checking and counting which words they use'. If someone says the word 'dog' three times in conversation, they are probably talking about dogs. This ignores many, many things about language and how it works, but is a good 'first guess' for what is being said. 

From our perspective, this means two things:
- We can broadly try to figure out the topic simply by checking which words are used
- We can choose a response based on a rule that defines what to do if and when we identify particular topics

In code then, this looks like:
- taking input as strings of continuous text, and splitting this up into lists that contain all the words we see
- use conditionals (`if`/`elif`/`else`) that choose an output based on what words appear in those lists

As we work through and create all the functions in the next section to create our chatbot, try to think through how these functions relate to this approach.

### Q9 - Is It a Question (0.5 points)

One thing we might want to know, given an input to a chatbot, is if the input is a question. 

We are going to use a very simple heuristic to check this: we will consider any input that has a question mark ('?') in it to be a question. 

To do so, write a function called `is_question`. 

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_string` - string

Output(s):
* `output` - boolean

Procedure(s):
- if there is a '?' in `input_string`, set `output` to `True`
- otherwise, set `output` to `False`. 
- `return` `output`

<!-- grade_id: cell-06800778424be5ba, points: 0, grade: false, solution: true -->
```python
def is_question(input_string):
    
    if ('?' in input_string) == True:
        output = True
    else:
        output = False
    
    return output

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-73224d2951e9985c, points: 0.25, grade: true, solution: false -->
```python
assert callable(is_question)
assert type(is_question('abc')) == bool
assert is_question('what?') == True
```

<!-- grade_id: cell-5b61947f11cc79af, points: 0.25, grade: true, solution: false -->
```python
assert callable(is_question)


```

### Example: string.punctuation

When working with strings, it can be useful to check what kind of things particular characters are. 

To do things like this, Python offers some collections of items that we can check against.

For example, we can ask if a particular character is a punctuation mark by asking if it is in the collection of punctuation characters. 

<span style="color:red">**Note**: If you get an error when trying to run the cell below, return to the top of the assignment and run the setup cell with the `import` statements.</span>


```python
# Print out list of punctuation
print(string.punctuation)

print() # Print blank line

# Check if a particular symbol is a punctuation mark
print('@' in string.punctuation)
```

### Q10 - Removing punctuation (0.5 points)

Once we know if an input is a question, we are going to get rid of all punctuation. 

To do so, write a function called `remove_punctuation`. 

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_string` - string

Output(s):
* `out_string` - string

Procedure(s):
- define a new variable `out_string` as an empty string
- loop through each character in `input_string`
    - if this character is not in `string.punctuation`
        - then concatenate the current character to `out_string`
- return `out_string`

<!-- grade_id: cell-e89fcaa0b5c38b5b, points: 0, grade: false, solution: true -->
```python
def remove_punctuation(input_string):
    
    out_string = ''
    for char in input_string:
        if char not in string.punctuation:
            out_string += char
            
    return out_string

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-7653654ef6ba58e1, points: 0.25, grade: true, solution: false -->
```python
assert callable(remove_punctuation)
assert type(remove_punctuation('a')) == str
assert remove_punctuation("Hey! It's Professor Morgan!") == "Hey Its Professor Morgan"
```

<!-- grade_id: cell-711049dd0e72709a, points: 0.25, grade: true, solution: false -->
```python
assert callable(remove_punctuation)


```

### Q11 - Preparing Text (0.5 points)

Now we want a more general function to prepare the text inputs for processing. 

Inside this function, we will use some string methods as well as our new `remove_punctuation` function.

Write a function called `prepare_text`.

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_string` - string

Output(s):
* `out_list` - list of strings

Procedure(s):
- make the string all lower case (use a string method called `lower`)
    - note that you will have to assign the output to a variable. You can use a variable name like `temp_string`
- remove all punctuation from the string (use your new function `remove_punctuation`)
    - this can also assign out to `temp_string`
- split the string into words (look for and use a string method called `split`)
    - note that split will return a list of strings. Assign this as `out_list`
- return `out_list`

<!-- grade_id: cell-c1c23cc422a6e514, points: 0, grade: false, solution: true -->
```python
def prepare_text(input_string):
    
    temp_string = input_string.lower()
    temp_string = remove_punctuation(temp_string)
    out_list = temp_string.split()
    
    return out_list

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-617b5f8e65de0792, points: 0.25, grade: true, solution: false -->
```python
assert callable(prepare_text)
assert type(prepare_text('One two three.')) == list
assert prepare_text('Hi! Also, howdy.') == ['hi', 'also', 'howdy']
```

<!-- grade_id: cell-9fb3a5fb59109d2c, points: 0.25, grade: true, solution: false -->
```python
assert callable(prepare_text)


```

### Q12 - Echo a Response (0.5 points)

We will define a function that will return a string that has been repeated a specified number of times, with a given separator. 

We will also want to deal with the case where the input to this function may be `None`. 

Write a function called `respond_echo`. 

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_string` - string, or None
* `number_of_echoes` - integer
* `spacer` - string

Output(s):
* `echo_output` - string, or None

Procedure(s):
- if `input_string` is not `None`:
    - set a variable `echo_output` to be  `input_string`, repeated the number of times specified by `number_of_echoes`, joined by the  `spacer`.
        - Hint: to do so, recall what "*" does to strings
- otherwise (if `input_string` is None):
    - set `echo_output` to `None`
- return `echo_output`

Note that a simple version of this function will end up adding the spacer to the end of the string as well. For our purposes here, implement that simple version.

So, for example, `respond_echo('bark', 2, '-')` should return: 'bark-bark-'

<!-- grade_id: cell-d77bc7ed64441ebc, points: 0, grade: false, solution: true -->
```python
def respond_echo(input_string, number_of_echoes, spacer):
    
    if input_string is not None:
        echo_output = (input_string+spacer)*(number_of_echoes)
    else:
        echo_output = None
    
    return echo_output    

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-3280829d10a28901, points: 0.25, grade: true, solution: false -->
```python
assert callable(respond_echo)
assert type(respond_echo('ha', 2, ' ')) == str
assert respond_echo('meow', 3, '~') == 'meow~meow~meow~'
assert respond_echo(None, 2, '') == None
```

<!-- grade_id: cell-e828c225b1081368, points: 0.25, grade: true, solution: false -->
```python
assert callable(respond_echo)


```

### Example: random.choice

We want our chatbot to have some variation in its responses - not always answering the same thing.

So, we want to be able to choose an output from a list of options. 

We can use the `random.choice` function to so.

An example of using `random.choice` is offered below.

```python
# Given a list of inputs, randomly choose and return one of them (re-run this to see different outputs)
random.choice(['first option', 'second option', 'third option'])
```

### Q13 - Selector (0.5 points)

Next we will write a function that will help us select an output for the chatbot, based on the input it got.

The overall goal of this function is to take a list of words that we got as input, a list of words to check for whether they appear in the input, and a list of possible outputs to return if something from the list to check is in the input list. 

Define a function called `selector`.

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_list` - list of strings
* `check_list` - list of strings
* `return_list` - list of strings

Output(s):
* `output` - string, or None

Procedure(s):
- Initialize `output` to `None`
- Loop through each element in `input_list`
- Use a conditional to check if the current element is in `check_list`
    - If it is, assign `output` as the returned value of calling the `random.choice` function on `return_list`
    - Also, break out of the loop
- At the end of the function, return `output`

Note that if we don't find any words from `input_list` that are in `check_list`, `output` will be returned as `None`.

<!-- grade_id: cell-a85973a26e490194, points: 0, grade: false, solution: true -->
```python
def selector(input_list, check_list, return_list):

    output = None
    for word in input_list:
        if word in check_list:
            output = random.choice(return_list)
            break

    return output

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-88e6e54e20d32647, points: 0.25, grade: true, solution: false -->
```python
assert callable(selector)

assert selector(['in', 'words'], ['words'], ['yes']) == 'yes'
assert selector(['in', 'words'], ['out'], ['yes']) == None
```

<!-- grade_id: cell-f38dab4bc4b62a22, points: 0.25, grade: true, solution: false -->
```python
assert callable(selector)


```

### Q14 - String Concatenator Function (0.5 points)

Create a function that will concatenate two strings, combining them with a specified separator. 

To do so, write a function called `string_concatenator`. 

This function should have the following inputs, outputs, and internal procedures:

Input(s)
* `string1` - string
* `string2` - string
* `separator` - string

Output(s)
* `output` - string

Procedure(s):
- concatenate `string1` to `string2` with `separator` in between them. 
- return the result

<!-- grade_id: cell-f24e8fdca335a582, points: 0, grade: false, solution: true -->
```python
def string_concatenator(string1, string2, separator):
    
    result = string1 + separator + string2
    
    return result

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-baa25e0ddd0ad3ac, points: 0.25, grade: true, solution: false -->
```python
assert callable(string_concatenator)
assert type(string_concatenator('hello', 'world', ' ')) == str
assert string_concatenator('hello', 'world', ' ') == 'hello world'
```

<!-- grade_id: cell-67395721ec848259, points: 0.25, grade: true, solution: false -->
```python
assert callable(string_concatenator)

assert string_concatenator('yay', 'python', ' ') == 'yay python'
assert string_concatenator('other', 'test', '-') == 'other-test'

```

### Q15 - List to String (0.5 points)

Since we are often dealing with a list of strings, it would be useful to have a function to turn a list of strings back into one single concatenated string. 

This function will return a string that is each element of `input_list` concatenated together with each separated by the string `separator`. 

For example, the following function call:

`list_to_string(['This', 'is', 'fun'], ' ')`

should return: 
    
    'This is fun'
    
To do this, write a function called `list_to_string`. 

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_list` - list of strings
* `separator` - string

Output(s):
* `output` - string

Procedure(s):
- assign a variable called `output` to be the first (index 0) element from `input_list`
- loop through the rest of `input_list`, looping from the 2nd element (index 1) through to the end of the list
    - Within the loop, use `string_concatenator` to combine `output` with the current element, separated by `separator`
        - Assign the output of this to be the new value of `output`
- return `output`, which should now be a list of the elements of `input_list` joined together into a single string

<!-- grade_id: cell-cfabdfe3966db8c6, points: 0, grade: false, solution: true -->
```python
def list_to_string(input_list, separator):
    
    output = input_list[0]
    for item in input_list[1:]:
        output = string_concatenator(output, item, separator)
        
    return output

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-b8d8c7506f53317e, points: 0.25, grade: true, solution: false -->
```python
assert callable(list_to_string)
assert type(list_to_string(['a', 'b'], '|')) == str
assert list_to_string(['a', 'b'], '|') == 'a|b'
```

<!-- grade_id: cell-e482a53f1553dc55, points: 0.25, grade: true, solution: false -->
```python
assert callable(list_to_string)


```

### Q16 - End Chat (0.5 points)

The last thing we need is a way to end the chat with our chatbot. 

We will use the same convention that `nltk` uses - namely, that we will end the conversation if the word `quit` appears.

To be able to do so, write a function called `end_chat`

This function should have the following inputs, outputs, and internal procedures:

Input(s):
* `input_list` - list

Output(s):
* `output` - boolean

Procedure(s):
- This function should return `True` if the string 'quit' is in `input_list`, and return `False` otherwise. 

<!-- grade_id: cell-1ab514a0dcaee24e, points: 0, grade: false, solution: true -->
```python
def end_chat(input_list):
    
    if 'quit' in input_list:
        output = True
    else:
        output = False
        
    return output

```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-7937207e3e365a07, points: 0.25, grade: true, solution: false -->
```python
assert callable(end_chat)
assert type(end_chat(['a', 'b'])) == bool
assert end_chat(['quit']) == True
```

<!-- grade_id: cell-d6260ce7e8bb5e3c, points: 0.25, grade: true, solution: false -->
```python
assert callable(end_chat)


```

## Pulling it all Together

If everything has gone well with writing everything above, then we are ready to pull all these functions together to create our chatbot.

The next few cells define a bunch of code, all provided for you. This code adds a couple more functions we can use, and then defines some lists of inputs and outputs that our chatbot will know about. After that, there is a big function that pulls everything together and creates the procedure to use all of this code together as a chatbot. 

You don't have to write anything in the cells below (or in the rest of the assignment). 

Do have a quick read through the code though, to see how our chatbot works, and to see how all the functions you just wrote are used to create the chatbot. 

If you need to, revisit the information at the beginning of this section about the logic our chatbot uses. 

You don't have to understand everything about what happens here, but try to have a first-pass look through it.

We will build up to understanding and writing this kind of code that brings a lot of things together as we continue through the course. 

If the code seems intimidating, just scroll down to the bottom and test out talking to our chatbot. 

```python
# These are a couple more functions, provided to you, to be used by our chatbot.

def is_in_list(list_one, list_two):
    """Check if any element of list_one is in list_two."""
    
    for element in list_one:
        if element in list_two:
            return True
    return False

def find_in_list(list_one, list_two):
    """Find and return an element from list_one that is in list_two, or None otherwise."""
    
    for element in list_one:
        if element in list_two:
            return element
    return None
```

```python
# This cell defines a collection of input and output things our chatbot can say and respond to

GREETINGS_IN = ['hello', 'hi', 'hey', 'hola', 'welcome', 'bonjour', 'greetings']
GREETINGS_OUT = ["Hello, it's nice to talk to you!", 'Nice to meet you!',  "Hey - Let's chat!"]

COMP_IN = ['python', 'code', 'computer', 'algorithm', ]
COMP_OUT = ["Python is what I'm made of.", \
            "Did you know I'm made of code!?", \
            "Computers are so magical", \
            "Do you think I'll pass the Turing test?"]

PEOPLE_IN = ['turing', 'hopper', 'neumann', 'lovelace']
PEOPLE_OUT = ['was awesome!', 'did so many important things!', 'is someone you should look up :).']
PEOPLE_NAMES = {'turing': 'Alan', 'hopper': 'Grace', 'neumann': 'John von', 'lovelace': 'Ada'}

JOKES_IN = ['funny', 'hilarious', 'ha', 'haha', 'hahaha', 'lol']
JOKES_OUT = ['ha', 'haha', 'lol'] 

NONO_IN = ['matlab', 'java', 'C++']
NONO_OUT = ["I'm sorry, I don't want to talk about"]

UNKNOWN = ['Good.', 'Okay', 'Huh?', 'Yeah!', 'Thanks!']

QUESTION = "I'm too shy to answer questions. What do you want to talk about?"
```

```python
def have_a_chat():
    """Main function to run our chatbot."""
    
    chat = True
    while chat:

        # Get a message from the user
        msg = input('INPUT :\t')
        out_msg = None

        # Check if the input is a question
        question = is_question(msg)

        # Prepare the input message
        msg = prepare_text(msg)

        # Check for an end msg 
        if end_chat(msg):
            out_msg = 'Bye!'
            chat = False

        # Check for a selection of topics that we have defined to respond to
        #   Here, we will check for a series of topics that we have designed to answer to
        if not out_msg:

            # Initialize to collect a list of possible outputs
            outs = []

            # Check if the input looks like a greeting, add a greeting output if so
            outs.append(selector(msg, GREETINGS_IN, GREETINGS_OUT))

            # Check if the input looks like a computer thing, add a computer output if so
            outs.append(selector(msg, COMP_IN, COMP_OUT))

            # Check if the input mentions a person that is specified, add a person output if so
            if is_in_list(msg, PEOPLE_IN):
                name = find_in_list(msg, PEOPLE_IN)
                outs.append(list_to_string([PEOPLE_NAMES[name], name.capitalize(),
                                            selector(msg, PEOPLE_IN, PEOPLE_OUT)], ' '))

            # Check if the input looks like a joke, add a repeat joke output if so
            outs.append(respond_echo(selector(msg, JOKES_IN, JOKES_OUT), 3, ''))

            # Check if the input has some words we don't want to talk about, say that, if so
            if is_in_list(msg, NONO_IN):
                outs.append(list_to_string([selector(msg, NONO_IN, NONO_OUT), find_in_list(msg, NONO_IN)], ' '))

            # IF YOU WANTED TO ADD MORE TOPICS TO RESPOND TO, YOU COULD ADD THEM IN HERE

            # We could have selected multiple outputs from the topic search above (if multiple return possible outputs)
            #   We also might have appended None in some cases, meaning we don't have a reply
            #   To deal with this, we are going to randomly select an output from the set of outputs that are not None
            options = list(filter(None, outs))
            if options:
                out_msg = random.choice(options)

        # If we don't have an output yet, but the input was a question, return msg related to it being a question
        if not out_msg and question:
            out_msg = QUESTION

        # Catch-all to say something if msg not caught & processed so far
        if not out_msg:
            out_msg = random.choice(UNKNOWN)

        print('OUTPUT:', out_msg)
```

## Talking to our Chatbot

To talk to our chatbot, uncomment and run the cell below. 

If anything seems to go weird, note that you can always interrupt the code from running the chat with kernel - interrupt from the menu. 

If this happens, it probably means there is some issue with one of the functions you wrote for the chatbot. 

```python
# Talk to our chatbot
# have_a_chat()
```

### Extending Our Chatbot

You might notice that our chatbot, right now, is pretty limited in how many things it can respond to.

To respond to more topics, we could re-use these same functions, adding more input & output options for our chatbot to respond to, using the basic approaches and functionality we have defined here. 

If you are interested in this topic, extending this chatbot (or ones like it) will be one of your project options. 

## The End!

After you've explored chatting with the chatbot, you are done!

Have a look back over your answers, and also make sure to Restart & Run All from the kernel menu to double check that everything is working properly and all of your asserts pass silently. 

When you are ready, submit on datahub.